Ultrastructural sublaminar-specific diversity of excitatory synaptic boutons in layer 1 of the adult human temporal lobe neocortex

Author response:

The following is the authors’ response to the previous reviews

Public Reviews:

Reviewer #1 (Public review):

Summary:

The Authors investigated the anatomical features of the excitatory synaptic boutons in layer 1 of the human temporal neocortex. They examined the size of the synapse, the macular or the perforated appearance and the size of the synaptic active zone, the number and volume of the mitochondria, the number of the synaptic and the dense core vesicles, also differentiating between the readily releasable, the recycling and the resting pool of synaptic vesicles. The coverage of the synapse by astrocytic processes was also assessed, and all the above parameters were compared to other layers of the human temporal neocortex. The Authors conclude that the subcellular morphology of the layer 1 synapses is suitable for the functions of the neocortical layer, i.e. the synaptic integration within the cortical column. The low glial coverage of the synapses might allow the glutamate spillover from the synapses enhancing synaptic crosstalk within this cortical layer.

Strengths:

The strengths of this paper are the abundant and very precious data about the fine structure of the human neocortical layer 1. Quantitative electron microscopy data (especially that derived from the human brain) are very valuable, since this is a highly time- and energy consuming work. The techniques used to obtain the data, as well as the analyses and the statistics performed by the Authors are all solid, strengthen this manuscript, and mainly support the conclusions drawn in the discussion.

Comments on latest version:

The corrected version of the article titled “Ultrastructural sublaminar specific diversity of excitatory synaptic boutons in layer 1 of the adult human temporal lobe neocortex" has been improved thanks to the comments and suggestions of the reviewers. The Authors implemented several of my comments and suggestions. However, many of them were not completed. It is understandable that the Authors did not start a whole new series of experiments investigating inhibitory synapses (as it was a misunderstanding affecting 2 reviewers from the three). But the English text is still very hard to understand and has many mistakes, although I suggested to extensively review the use of English. Furthermore, my suggestion about avoiding many abbreviations in the abstract, analyse and discuss more the perforated synapses, the figure presentation (Figure 3) and including data about the astrocytic coverage in the Results section were not implemented. My questions about the number of docked vesicles and p10 vesicles, as well as about the different categories of the vesicle pools have not been answered neither. Many other minor comments and suggestions were answered, corrected and implemented, but I think it could have been improved more if the Authors take into account all of the reviewers' suggestions, not only some of them. I still have several main and minor concerns, with a few new ones as well I did not realize earlier, but still think it is important.

We would like to thank the reviewer for the comments.

- We worked on the English again and tried to improve the language.

- We avoided to use too many abbreviations in the Abstract and reduced them to a minimum.

- We included a small paragraph about non-perforated vs. perforated active zones in both the Results and Discussion sections. However, since the majority of active zones in all cortical layers of the human TLN were of the macular type, we concluded that it is not relevant to describe their function in more detail.

- In Figure 3 A-C we added contour lines to the boutons to make their outlines more visible.

- We completed the data about the astrocytic coverage in the Results section (see also below).

- Concerning the vesicle pools please see below.

Main concerns:

(1) Epileptic patients:

As all patients were epileptic, it is not correct to state in the abstract that non-epileptic tissue was investigated. Even if the seizure onset zone was not in the region investigated, seizures usually invade the temporal lobe in TLE. If you can prove that no spiking activity occurred in the sample you investigated and the seizures did not invade that region, then you can write that it is presumably non-epileptic. I would suggest to write “L1 of the human temporal lobe neocortical biopsy tissue". See also Methods lines 608-612. Write only “non-epileptic" or “non-affected" if you verified it with EcoG. If this was the case, please write a few sentences about it in the Methods.

We rephrased Material and Methods concerning this point and added that patients were monitored with EEG, MRI and multielectrode recordings. In addition, we stated that the epileptic focus was always far away from the neocortical tissue samples. Furthermore, we added a small paragraph that functional studies using the same methodology have shown that neocortical access tissue samples taken from epilepsy surgery do not differ in electrophysiological properties and synaptic physiology when compared with acute slice preparations in experimental animals and we quoted the relevant papers.

We hope that the reviewer is now convinced that our tissue samples can be regarded as non-affected.

(2) About the inhibitory/excitatory synapses.

Since our focus was on excitatory synaptic boutons as already stated in the title we have not analyzed inhibitory SBs. Now, I do understand that only excitatory synapses were investigated. Although it was written in the title, I did not realized, since all over the manuscript the Authors were writing synapses, and were distinguishing between inhibitory and excitatory synapses in the text and showing numerous excitatory and inhibitory synapses on Figure 2 and discussing inhibitory interneurons in the Discussion as well. Maybe this was the reason why two reviewers out of the three (including myself) thought you investigated both types of synapses but did not differentiated between them. So, please, emphasize in the Abstract (line 40), Introduction (for ex. line 92-97) and the Discussion (line 369) that only excitatory synaptic boutons were investigated.

As this paper investigated only excitatory synaptic boutons, I think it is irrelevant to write such a long section in the Discussion about inhibitory interneurons and their functions in the L1 of the human temporal lobe neocortex. Same applies to the schematic drawing of the possible wiring of L1 (Figure 7). As no inhibitory interneurons were examined, neither the connection of the different excitatory cells, only the morphology of single synaptic boutons without any reference on their origin, I think this figure does not illustrate the work done in this paper. This could be a figure of a review paper about the human L1, but is inappropriate in this study.

We followed the reviewer’s suggestion and pointed out explicitly that we only investigated excitatory synaptic boutons. We also changed the Discussion and focused more on circuitry in L1 and the role of CR-cells.

(3) Perforated synapses

The findings of the Geinismann group suggesting that perforated synapses are more efficient than non-perforated ones is nowadays very controversially discussed” I did not ask the Authors to say that perforated synapses are more efficient. However, based on the literature (for ex. Harris et al, 1992; Carlin and Siekievitz, 1982; Nieto-Sampedro et al., 1982) the presence of perforated synapses is indeed a good sign of synapse division/formation - which in turn might be coupled to synaptic plasticity (Geinisman et al, 1993), increased synaptic activity (Vrensen and Cardozo, 1981), LTP (Geinisman et al, 1991, Harris et al, 2003), pathological axonal sprouting (Frotscher et al, 2006), etc. I think it is worth mentioning this at least in the Discussion.

We agree with the reviewer and added a small paragraph in the Results section about the two types of AZs in L1 of the human TLN. We pointed out that there are both types, macular non-perforated and perforated AZs, but the majority in all layers were of the non-perforated type. In the Discussion we added some paper pointing out the role of perforated synapses.

(4) Question about the vesicle pools

Results, Line 271: Still not understandable, why the RRP was defined as {less than or equal to}10 nm and {less than or equal to}20nm. Why did you use two categories? One would be sufficient (for example {less than or equal to}20nm). Or the vesicles between 10 and 20nm were considered to be part of RRP? In this case there is a typo, it should be {greater than or equal to}10 nm and {less than or equal to}20nm.

The answer of the Authors was to my question raised: We decided that also those very close within 10 and 20 nm away from the PreAZ, which is less than a SV diameter may also contribute to the RRP since it was shown that SVs are quite mobile.

This does not clarify why did you use two categories. Furthermore, I did not receive answer (such as Referee #2) for my question on how could you have 3x as many docked vesicles than vesicles {less than or equal to}10nm. The category {less than or equal to}10nm should also contain the docked vesicles. Or if this is not the case, please, clarify better what were your categories.

We thank the reviewer for pointing out that mentioning two distance criteria (p10 and p20) to define one physiological entity (RRP) is somewhat confusing and we acknowledge that the initial response to the reviewers falls short of explaining this choice. This is indeed only understandable in the context of the original paper by Sätzler et al. 2002, where these criteria were first introduced. We therefore referenced this publication more prominently in the paragraph in question.

So to explain this, we first would like to clarify the definition of the two RRP classification criteria used (p10 and p20), which has caused some confusion amongst the reviewers as to which vesicles where included or not:

- p10 criterion: p£10 nm (SVs have a minimum distance less than or equal to 10 nm from the PreAZ), including ‘docked’ vesicles which have a distance of zero or less (p0)

- p20 criterion: p£20 nm (SVs have a minimum distance less than or equal to 20 nm from the PreAZ), including vesicles of the p10 criterion.

As mentioned, these criteria were introduced first in Sätzler et al. 2002 looking at the Calyx of Held synapse. In that paper, we tried to establish a morphological correlate to existing physiological measurements, which included the RRP. As there is no known marker that would allow to discriminate between vesicles that contribute to the RRP anatomically, we looked at existing physiological experiments such as Schneggenburger et al. 1999; Wu and Borst 1999; Sun and Wu 2001 and compared their total numbers to our measurements. As the number of docked vesicles (p0, see above) was on the lower side of these physiological estimates, we also looked at vesicles close to the AZ, which we think could be recruited within a short time (£ 10 msec). Comparing with existing literature, we found that at p20 we get pool sizes comparable to midrange estimates of reported RRP sizes. In order to account for the variability of the observed physiological pool sizes, we reported all three measurements (p0, p10, p20) not only in the original Calyx of Held, but in all subsequent studies of different CNS synapses of our group since then.

As it remains uncertain if such correlate indeed exists, we therefore followed the suggestion to rephrase RRP and RP to putative RRP and putative RP (see also Rollenhagen et al. 2007). We thank both reviewers for pointing out this omission.

Concerning the difference between ‘docked’ vesicles and vesicles within the p10 perimeter criterion. First of all, the reviewer is right in saying that the category p10 ({less than or equal to}10nm) should also contain the docked vesicles (see above). The fact to have 3x as many ‘docked’ vesicles in our TEM tomography than in the p10 distance analysis could be partly explained, on the one hand, by a very high variability between patients (as expressed by the high SD, table 1) and, on the other hand, by a high intraindividual synaptic bouton variability. In both sublayers, there is a huge difference in the number of vesicles within the p10 criterion of individual synaptic boutons ranging from 0 to ~40 with a mean value of ~1 to ~4 (calculated per patient), the upper level being close to the values calculated with TEM tomography for the ‘docked’ vesicles.

(5) Astrocytic coverage

On Fig. 6 data are presented on the astrocytic coverage derived from L1 and L4. In my previous review I asked to include this in the text of the Results as well, but I still do not see it. It is also lacking from the Results how many samples from which layer were investigated in this analysis. Only percentages are given, and only for L1 (but how many patients, L1a and/or L1b and/or L4 is not provided). In contrast, Figure 6 and Supplementary Table 2 (patient table) contains the information that this analysis has been made in L4 as well. Please, include this information in the text as well (around lines 348-360).

In our previous revised version, we had included the values shown in Fig. 6 for both L1 and L4 in the Results section (L4: lines 352 – 355: ‘The findings in L1…’). However, we agree with the reviewer and have now also added the number of patients and synapses investigated (now lines 359 – 365).

About how to determine glial elements. I cannot agree with the Authors that glial elements can be determined with high certainty based only on the anatomical features of the profiles seen in the EM. “With 25 years of experience in (serial) EM work" I would say, that glial elements can be very similar to spine necks and axonal profiles.

All in all, if similar methods were used to determine the glial coverage in the different layers of the human neocortex, than it can be compared (I guess this is the case). However, I would say in the text that proper determination would need immunostaining and a new analysis. This only gives an estimation with the possibility of a certain degree of error.

We do not entirely agree with the reviewer on this point. As stated in the text, there are structural criteria to identify astrocytic elements (see citations quoted). These golden standard criteria are commonly used also by other well-known groups (DeFelipe and co-workers, Francisco Clasca and co-workers; Michael Frotscher the late and co-workers etc.). However, in a past paper about astrocytic coverage of synaptic complexes in L5 of the human TLN, immunohistochemistry against glutamine synthetase, a key enzyme in astrocytes, was carried out to describe the coverage. This experiment supports our findings in the other cortical layers of the human TLN. As the reviewer might know, immunohistochemistry always led to a reduction in ultrastructural preservation, so we decided not to use immunohistochemistry for the further publications of the other cortical layers. We added a short notice on this in the Material and Methods section.

(6) Large interindividual differences in the synapse density should be discussed in the Discussion.

As suggested by the reviewer we have included a sentence in the Discussion that interindividual differences can be either related to differences in age, gender and the use of different methodology as suggested by DeFelipe and co-workers (1999)

Reviewer #2 (Public review):

Summary:

The study of Rollenhagen et al examines the ultrastructural features of Layer 1 of human temporal cortex. The tissue was derived from drug-resistant epileptic patients undergoing surgery, and was selected as further from the epilepsy focus, and as such considered to be non-epileptic. The analyses has included 4 patients with different age, sex, medication and onset of epilepsy. The MS is a follow-on study with 3 previous publications from the same authors on different layers of the temporal cortex:

Layer 4 - Yakoubi et al 2019 eLife

Layer 5 - Yakoubi et al 2019 Cerebral Cortex,

Layer 6 - Schmuhl-Giesen et al 2022 Cerebral Cortex

They find, the L1 synaptic boutons mainly have single active zone a very large pool of synaptic vesicles and are mostly devoid of astrocytic coverage.

Strengths:

The MS is well written easy to read. Result section gives a detailed set of figures showing many morphological parameters of synaptic boutons and surrounding glial elements. The authors provide comparative data of all the layers examined by them so far in the Discussion. Given that anatomical data in human brain are still very limited, the current MS has substantial relevance.

The work appears to be generally well done, the EM and EM tomography images are of very good quality. The analyses is clear and precise.

Weaknesses:

The authors made all the corrections required, answered most of my concerns, included additional data sets, and clarified statements where needed.

My remaining points are:

Synaptic vesicle diameter (that has been established to be ~40nm independent of species) can properly be measured with EM tomography only, as it provides the possibility to find the largest diameter of every given vesicle. Measuring it in 50 nm thick sections result in underestimation (just like here the values are ~25 nm) as the measured diameter will be smaller than the true diameter if the vesicle is not cut in the middle, (which is the least probable scenario). The authors have the EM tomography data set for measuring the vesicle diameter properly.

We thank the reviewer for the helpful comments. We followed the recommendation to measure the vesicle diameter using our TEM tomography tilt series, but came to similar results concerning this synaptic parameter. As stated in our Material and Methods section, we only counted (measured) clear ring-link structures according to a paper by Abercrombie (1963). Since our results are similar for both methods, we do believe that our measurements are correct. Even random single measurements on the original 3D tilt-series yielded comparable results (Lübke and co-workers, personal observation). Furthermore, our results are within ranges, although with high variability, also described by other groups (see discussion lines 436 - 449). We therefore hope that the reviewer will now accept our measurements.

It is a bit misleading to call vesicle populations at certain arbitrary distances from the presynaptic active zone as readily releasable pool, recycling pool and resting pool, as these are functional categories, and cannot directly be translated to vesicles at certain distances. Even it is debated whether the morphologically docked vesicles are the ones, that are readily releasable, as further molecular steps, such as proper priming is also a prerequisite for release.

It would help to call these pools as "putative" correlates of the morphological categories.

We followed the suggestion by the reviewer and renamed our vesicle pools as putative RRP, putative RP and putative resting pools.

Reviewer #3 (Public review):

Summary:

Rollenhagen at al. offer a detailed description of layer 1 of the human neocortex. They use electron microscopy to assess the morphological parameters of presynaptic terminals, active zones, vesicle density/distribution, mitochondrial morphology and astrocytic coverage. The data is collected from tissue from four patients undergoing epilepsy surgery. As the epileptic focus was localized in all patients to the hippocampus, the tissue examined in this manuscript is considered non-epileptic (access) tissue.

Strengths:

The quality of the electron microscopic images is very high, and the data is analyzed carefully. Data from human tissue is always precious and the authors here provide a detailed analysis using adequate approaches, and the data is clearly presented.

Weaknesses:

The text connects functional and morphological characteristics in a very direct way. For example, connecting plasticity to any measurement the authors present would be rather difficult without any additional functional experiments. References to various vesicle pools based on the location of the vesicles is also more complex than it is suggested in the manuscript. The text should better reflect the limitations of the conclusions that can be drawn from the authors' data.

Recommendations for the authors:

Reviewer #1 (Recommendations for the authors):

Astrocytic coverage

On Fig. 6 data are presented on the astrocytic coverage derived from L1 and L4. In my previous review I asked to include this in the text of the Results as well, but I still do not see it. It is also lacking from the Results how many samples from which layer were investigated in this analysis. Only percentages are given, and only for L1 (but how many patients, L1a and/or L1b and/or L4 is not provided). In contrast, Figure 6 and Supplementary Table 2 (patient table) contains the information that this analysis has been made in L4 as well. Please, include this information in the text as well (around lines 348-360).

See above.

About how to determine glial elements. I cannot agree with the Authors that glial elements can be determined with high certainty based only on the anatomical features of the profiles seen in the EM. “With 25 years of experience in (serial) EM work" I would say, that glial elements can be very similar to spine necks and axonal profiles. Please, see the photos below, out of the 16 circled profiles (2nd picture, very similar to each other) only 3 belong to an astroglial cell (last picture, purple profiles-purple cell), 10 are spines/spine necks/small caliber dendrites of pyramidal cells, 3 are axonal profiles (last but one picture, blue profiles, marked with arrows on the right side). If you follow in your serial sections those elements which you think are glial processes and indeed they are attached to a confidently identifiable glial cell, I agree, it is a glial process. But identifying small, almost empty profiles without any specific staining, from one single EM section, as glial process is very uncertain. Please, check the database of the Allen Institute made from the V1 visual cortex of a mouse. It is a large series of EM sections where they reconstructed thousands of neurons, astroglial and microglial cells. It is possible to double click on the EM picture on a profile and it will show the cell to which that profile belongs. https://portal.brain-map.org/connectivity/ultrastructural-connectomics Pictures included here: https://elife-rp.msubmit.net/eliferp_files/2024/11/25/00132644/02/132644_2_attach_21_29456_convrt.pdf

All in all, if similar methods were used to determine the glial coverage in the different layers of the human neocortex, than it can be compared (I guess this is the case). However, I would say in the text that proper determination would need immunostaining and a new analysis. This only gives an estimation with the possibility of a certain degree of error.

As stated above, we carried out glutamine synthetase immunohistochemistry in L5 of the human TLN and came to the same results. However, we added a sentence on this in the chapter on astrocytic coverage in the Material and Methods section. Additionally, we modified this chapter according to the reviewer’s suggestion.

Minor comments

Introduction: Last sentence is not understandable (lines 101-103), please rephrase. (contribute to understand or contribute in understanding or contribute to the understanding of..., but definitely not contribute to understanding). The authors should check and review extensively for improvements to the use of English, or use a program such as Grammarly.

Results: Grammar (line 107): L1 in the adult mammalian neocortex represents a relatively...

Line 173: “Some SBs in both sublaminae were seen to establish either two or three SBs on the same spine, spines 173 of other origin or dendritic shafts." - Some SBs established two or three SBs? I would write Some SBs established two or three synapses on...

Line 243: “The synaptic cleft size were slightly, but non-significantly different"

Line 260: “DCVs play an important role in endo- and exocytosis, the build-up of PreAZs by releasing Piccolo and Bassoon (Schoch and Gundelfinger 2006; Murkherjee et al. 2010)," - please, correct this.

We have done corrections as suggested by the reviewer.

Line 374: No point at the end of the last phrase.

Discussion:

Lines 400-404: “The majority of SBs in L1 of the human TLN had a single at most three AZs that could be of the non perforated macular or perforated type comparable with results for other layers in the human TLN but by ~1.5-fold larger than in rodent and non-human primates." - What is comparable with the other layers, but different from animals? Please rephrase this sentence, it is not understandable. I already mentioned this sentence in my previous review, but nothing happened.

Lines 435-437: “Remarkably, the total pool sizes in the human TLN were significantly larger by more than 6-fold (~550 SVs/AZ), and ~4.7-fold (~750 SVs/AZ;) than those in L4 and L5 (Yakoubi et al. 2019a, b; see also Rollenhagen et al. 2018) in rats." Please rethink what you wished to say and compare to the sentence meaning. I think you wanted to compare human TLN L1 pool size to L4 and L5 in the human TLN (Yakoubi 2019a and b) and to rat (Rollenhagen 2018). Instead, you compared all layers of the human TLN to L4 and L5 in rats (with partly wrong references). Please rephrase this. Lines 483-484: “Astrocytes serve as both a physical barrier to glutamate diffusion and as mediate neurotransmitter uptake via transporters".

This sentence is grammatically incorrect, please rephrase.

We corrected the sentences as suggested by the reviewer.

Methods:

In the text, there are only 4 patients (lines 603-604), but in the supplementary table there are 9 patients (5 new included for L4 astrocytic coverage). Please, correct it in the text.

Lines 608-609: “neocortical access tissue samples were resected to control the seizures for histological inspection by neuropathologists." - What is the meaning of this? Please, rephrase.

We thank the reviewer for the comment and included the 5 patients used for L4 to the Material and Methods section, as well as in the Results section.

The reviewer is right, and we rephrased and corrected the sentence concerning the inspection by neuropathologists.

Figures

Figures 5B: The legend says “SB (sb) synapsing on a stubby spine (sp) with a prominent spine apparatus (framed area) and a thick dendritic segment (de) in L1b" - In my opinion this is not one synaptic bouton, but two. Clearly visible membranes separate them, close to the spine.

Supplemental Table 2 (patient table). If there is no information about Hu_04 patient's epilepsy, please write N/A (=non available) instead of - (which means it does not exist).

The reviewer is right, and we corrected the figure and the legend, as well as the table accordingly.

Reviewer #2 (Recommendations for the authors):

The authors addressed almost all of my concern, only this one remained:

If there is, however, relevant literature on "methods based on EM tomography" and "stereological methods to estimate both types of error" (over- and underestimates) that we are missing out on, we would appreciate the reviewer providing us with the corresponding references so that we can include such calculations in our paper.

There is a very detailed new study on calculating correction for TEM 2D 3D, Rothman et al 2023 PLOS One. That addresses most of these issues.

We thank the reviewer for drawing our attention to the publication by Rothman et al. 2023, which is a very detailed and comprehensive study looking at accurately estimating distributions of 3D size and densities of particles from 2D measurements using – amongst others – ET and TEM images as well as synaptic vesicles for validating their method. However, we do not see how this would be relevant to the reported mean diameters and their corresponding variances. And even if we would have reported on vesicle size/diameter distributions (referred to as G(d) in Rothmann et al. 2023), the authors themselves state that “… the results from our ET and TEM image analysis highlight the difficulty in computing a complete G(d) of MFT vesicles due to their small size…

Changes from the earlier version to the revised version

Abstract

Abstract

Earlier version: Layer (L)1, beside receiving massive cortico-cortical, commissural and associational projections, is the termination zone of tufted dendrites of pyramidal neurons and the area of Ca2+ spike initiation. However, its synaptic organization in humans is not known. Quantitative 3D-models of synaptic boutons (SBs) in L1 of the human temporal lobe neocortex were generated from non-epileptic neocortical biopsy tissue using transmission electron microscopy, 3D-volume reconstructions and TEM tomography. Particularly, the size of active zones (AZs) and the readily releasable, recycling and resting pool of synaptic vesicles (SVs) were quantified.

Revised version: Layer (L)1, beside receiving massive cortico-cortical, commissural and associational projections, is the termination zone of tufted dendrites of pyramidal neurons and the area of Ca2+ spike initiation. However, its synaptic organization in humans is not known. Quantitative 3D-models of excitatory synaptic boutons (SBs) in layer 1 of the human temporal lobe neocortex were generated from non-affected neocortical biopsy tissue using transmission electron microscopy, 3D-volume reconstructions and TEM tomography. Particularly, the size of active zones (AZs) and the readily releasable, recycling and resting pool of synaptic vesicles (SVs) were quantified.

Earlier version: The majority of SBs contained numerous mitochondria comprising ∼7% of the total volume of SBs, had a large macular, non-perforated AZs (∼0.20 µm2) and were predominantly located on dendritic spines. SBs had a total pool of ∼3500 SVs, a comparably large readily releasable (∼4 SVs), recycling (∼470 SVs) and resting (∼2900 SVs) pool. Astrocytic coverage of SBs suggests both synaptic cross talk or removal of spilled glutamate by astrocytic processes at synaptic complexes.

Revised version: The majority of excitatory SBs contained numerous mitochondria comprising ∼7% of the total volume, had a large macular, non-perforated AZ (∼0.20 µm2) and were predominantly located on dendritic spines. Excitatory SBs had a total pool of ∼3500 SVs, a comparably large readily releasable (∼4 SVs), recycling (∼470 SVs) and resting (∼2900 SVs) pool. Astrocytic coverage of excitatory SBs suggests both synaptic cross talk or removal of spilled glutamate by astrocytic processes at synaptic complexes.

The structural composition of SBs in L1 may underlie the function of L1 networks that mediate, integrate and synchronize contextual and cross-modal information, enabling flexible and state-dependent processing of feedforward sensory inputs from other layers of the cortical column.

Introduction

The mammalian neocortex is the most complex part of the brain comprising over 75% of the gray matter. One fundamental feature established during development is the formation of distinct layers with different networks of neurons, and its organization into vertically oriented slabs, so-called cortical columns (Marin-Padilla 1978, 1998; reviewed by Rockland and DeFelipe 2018). This organization is established in an ‘inside layers first-outside layers last’ fashion (Luskin and Shatz 1985a, b; reviewed by Cooper et al. 2008). The exception, beside layer (L)6 is L1 generated first originating from the marginal zone (Marin Padilla 1998; reviewed by Bystron et al. 2008). Ongoing cortico- and synaptogenesis finally lead to the formation of the adult cortical network with its unique structural and functional properties (reviewed by Lübke and Feldmeyer 2007; Rockland and DeFelipe 2018).

Earlier version: Remarkably, L1 is highly conserved across cortical areas and higher mammalian species, but its importance was long underestimated. However, L1 is now considered as the predominant input layer for top-down information, relayed by a rich, dense network of cortico-cortical, cortico-thalamic, commissural and associational long-range axonal projections. In humans, and thus different to experimental animals, both L1 excitatory Cajal-Retzius (CR) cells (Anstötz et al. 2014) and a quite heterogeneous population of inhibitory neurons (see for example Jiang et al. 2013; Boldog et al. 2018; Obermayer et al. 2018; Kwon et al. 2019; reviewed by Huang et al. 2024) providing signals to the terminal tuft dendrites of pyramidal neurons of the underlying cortical layers. Thus, L1 is a central locus of neocortical associations controlled by distinct types of inhibition and feedforward excitation (Schumann et al. 2019; Hartung and Letzkus 2021, reviewed by Schumann et al. 2021).

Revised version: Remarkably, L1 is highly conserved across cortical areas and higher mammalian species, but its importance was long underestimated. However, L1 is now considered as the predominant input layer for top-down information, relayed by a rich, dense network of cortico-cortical, cortico-thalamic, commissural and associational long-range axonal projections. In humans, and thus different to experimental animals, both L1 excitatory Cajal-Retzius (CR) cells (Anstötz et al. 2014) and a quite heterogeneous population of inhibitory neurons (see for example Jiang et al. 2013; Boldog et al. 2018; Obermayer et al. 2018; Kwon et al. 2019; reviewed by Huang et al. 2024) provide signals to the terminal tuft dendrites of pyramidal neurons of the underlying cortical layers. Thus, L1 is a central locus of neocortical associations controlled by distinct types of inhibition and feedforward excitation (Schumann et al. 2019; Hartung and Letzkus 2021, reviewed by Schumann et al. 2021).

Earlier version: The temporal lobe neocortex (TLN) is situated on the basolateral aspect of the cerebral hemispheres primarily observed in higher primates, including humans. Notably, in human, the TLN encompasses approximately 20% of the entire cerebral cortex (Kiernan 2012). It is recognized as a highly specialized associative neocortex, characterized by its homotypic granular six-layered organization (von Economo and Koskinas 1925; Palomero-Gallagher and Zilles 2019). The human TLN plays a crucial role in various cognitive functions, including auditory, visual, vestibular, linguistic, and olfactory processing. Additionally, it is intricately connected to diverse sensory and multimodal associational brain regions, such as the limbic system, amygdala, and various subcortical structures (Insausti 2013). Lastly, the human TLN serves as the epicenter and onset site for temporal lobe epilepsy (Allone et al. 2017; Tai et al. 2018).

Revised version: The temporal lobe neocortex (TLN) is situated on the basolateral aspect of the cerebral hemispheres primarily observed in higher primates, including humans. Notably, in humans, the TLN encompasses approximately 20% of the entire cerebral cortex (Kiernan 2012). It is recognized as a highly specialized associative neocortex, characterized by its homotypic granular six-layered organization (von Economo and Koskinas 1925; Palomero-Gallagher and Zilles 2019). The human TLN plays a crucial role in various cognitive functions, including auditory, visual, vestibular, linguistic, and olfactory processing. Additionally, it is intricately connected to diverse sensory and multimodal associational brain regions, such as the limbic system, amygdala, and various subcortical structures (Insausti 2013). Lastly, the human TLN serves as the epicenter and onset site for temporal lobe epilepsy (Allone et al. 2017; Tai et al. 2018).

One of the most important discoveries in neuroscience, beside the definition of the neuron by Cajal’s neuronal doctrine (1911), was the introduction of the term ‘Synapse’, more than 100 years ago. Since then these structures have been investigated from different viewpoints including structural, functional, molecular and computational studies summarized in meanwhile thousands of original publications, reviews and numerous textbooks. Yet, the quantitative geometry of the most prevalent synapse in the brain, the cortical synapse, has remained largely unexplored structurally, in both experimental animals (but see Rollenhagen et al. 2015, 2018; Bopp et al. 2017; Hsu et al. 2017; Prume et al. 2020) and, even more so, in the context of the human brain (but see Yakoubi et al. 2019a, b; Dominguez-Alvaro et al. 2019, 2021; Cano-Astorga et al. 2021, 2023; Schmuhl-Giesen et al. 2022; Shapson-Coe et al. 2024). This could be partially attributed to the availability of human brain tissue samples. Nevertheless, detailed and quantitative descriptions of synaptic complexes in L1 of the human brain are to date non-existing.

Earlier version: Here, an investigation into the synaptic organization of L1 in the human TLN was undertaken using high-resolution transmission electron microscopy (TEM), 3D-volume reconstruction to generate quantitative 3D-models of synaptic boutons (SBs) and by TEM tomography. The ultimate objective of this comprehensive project is to elucidate the synaptic organization, layer by layer, of the cortical column in humans.

Revised version: Here, an investigation into the synaptic organization of L1 in the human TLN was undertaken using high-resolution transmission electron microscopy (TEM), 3D-volume reconstruction, to generate quantitative 3D-models of synaptic boutons (SBs), and TEM tomography. The ultimate objective of this comprehensive project is to elucidate the synaptic organization, layer by layer, of the cortical column in humans.

Earlier version: Such meticulous quantitative structural analyses of SBs and their corresponding target structures are indispensable for comprehending, elucidating, and establishing connections between the structural and functional properties of the diverse signal cascades underlying synaptic transmission, efficacy, strength, and plasticity. Consequently, these efforts contribute significantly to understanding the computational properties of brain networks, with the organization of the cortical column emerging as a pivotal component thereof.

Revised version: Such meticulous quantitative structural analyses of SBs and their corresponding target structures are indispensable for comprehending, elucidating, and establishing connections between the structural and functional properties of the diverse signal cascades underlying synaptic transmission, efficacy, strength, and plasticity. Consequently, these efforts contribute significantly to the understanding of the computational properties of brain networks, in particular the synaptic organization of the cortical column.

Results

Neural and synaptic composition of L1 in the human TLN

Earlier version: L1 in the adult mammalian neocortex represent a relatively cell sparse layer that in humans can be subdivided into two sublaminae, L1a and L1b as revealed in semithin sections (Figure 1A) and Golgi-preparations (Figure 1B, C, F, G) and at the EM level (Figure 1H, I). L1a is more dominated by a dense network of astrocytes of different shape and size, their fine processes (Figure 2A), the occurrence of reactive microglia (Figures 1H, I, 2A) and terminal tuft dendrites of pyramidal neurons throughout the neuropil of L1 (Figures 1H, I, 2B).

Revised version: L1 in the adult mammalian neocortex represents a relatively cell sparse layer, that in humans can be subdivided into two sublaminae, L1a and L1b, as revealed in semithin sections (Figure 1A) and Golgi-preparations (Figure 1B, C, F, G), as well as at the EM level (Figure 1H, I). L1a is more dominated by a dense network of astrocytes of different shape and size, their fine processes (Figure 2A), the occurrence of reactive microglia (Figures 1H, I, 2A) and terminal tuft dendrites of pyramidal neurons throughout the neuropil of L1 (Figures 1H, I, 2B).

L1 and L2 in the human TLN as revealed in methylene-blue stained semithin sections, Golgi-impregnation and low power EM.

A, Methylene-blue stained semithin section of L1a, L1b, and L2. The cell sparse zone contains degenerating neurons (dark appearance, arrowheads in the framed area) and a cell body (asterisk in the framed area). Note the sudden increase in the density of neurons at the L1b/L2 border. Scale bar 500 µm. B, C, Golgi-impregnated radial astrocytes in L1a (B) and L1b (C). Scale bars 50 µm. D, Group of neurons in L1b containing lipofuscin granules of various shape and size (asterisks) most prominent in two of the neurons. Scale bar 50 µm. E, Putative CR-cell identifiable by its long horizontally oriented dendrite (marked by arrowheads) and an inverted putative GABAergic interneuron (asterisk) in L1b. Scale bar 50 µm. F, G, Golgi-impregnated CR-Cells in L1a with the characteristic horizontal orientation of dendrites (arrowheads) and axons (asterisks). Scale bars 50 µm. H, EM micrograph of L1a underneath the pial surface containing two putative CR-cells (asterisks). Note the abrupt transition (indicated by the dashed line) to the neuropil of L1b, and a reactive microglial cell (arrowhead), and dendrites and synapses. Scale bar 25 µm. I, EM micrograph of L1b containing thousands of dendritic profiles, a putative GABAergic interneuron (asterisk) but also degenerating neurons (arrowheads) and a reactive microglia (I1) Scale bar 50 µm.

Structural and synaptic organization of L1 in the human TLN as revealed by EM.

A, Superficial part of L1a composed of a dense network of fine astrocytic processes (asterisks). In the reactive microglial cell (rmic) so-called lipofuscin granules (framed area) were observed. Below the astrocytic network synaptic complexes gradually increased (highlighted in blue for target structures, yellow for SBs). Scale bar 10 µm. B, Two apical dendrites (blue) traversing L2 with beginning terminal tuft dendrites in L2 and L1 contacted by SBs (transparent yellow); some of which are putatively GABAergic terminating on the dendritic shaft (asterisks). Scale bar 2.5 µm. C, Large GABAergic SB (sb, transparent yellow) identifiable by the small-sized ovoid SVs and the less prominent PSD (arrowheads) synapsing on a dendritic shaft (de, transparent blue) in L1b. D, Three SBs in L1a (sb1-sb3, transparent yellow) located on a dendritic spine (sp, transparent blue) and a neighboring dendritic shaft (sh, transparent blue). Note that sb1 is excitatory as indicated by the larger round SVs with a prominent AZ whereas sb2 and sb3 are putative inhibitory terminals identified by smaller ovoid SVs and the lack of a prominent PSD. E, Elongated spine (sp) emerging from a small caliber dendrite (de, transparent blue) receiving input from a small-sized SB (sb, transparent yellow). Note the large, elongated spine apparatus (framed area). Scale bar in C-E 0.5 µm. In graphs C-E AZs are marked by arrowheads.

Earlier version: Hence, L1a can be described as a sublamina with a predominant proportion of astrocytic and microglial processes, although it also harbors numerous other structural elements such as terminal tuft dendrites and synaptic complexes. These complexes consist of a SB juxtaposed with either a dendritic shaft or spine (Figure 2A) directly beneath or interwoven among the fine astrocytic processes within L1 (Figure 2A). Lipofuscin granules, indicative of aging, were frequently noted in astrocytes, reactive microglia (Figure 2A) and L1 neurons, although their abundance, size, and morphology varied considerably.

Revised version: Hence, L1a can be described as a sublamina with a predominant proportion of astrocytes, microglia and their processes, although it also harbors numerous other structural elements such as terminal tuft dendrites and synaptic complexes. These complexes consist of a SB juxtaposed with either a dendritic shaft or spine (Figure 2A) directly beneath or interwoven among the fine astrocytic processes within L1 (Figure 2A). Lipofuscin granules, indicative of aging, were frequently noted in astrocytes, reactive microglia (Figure 2A) and L1 neurons, although their abundance, size, and morphology varied considerably.

Earlier version: Remarkably, and in contrast to rodents a persistent subpopulation of excitatory CR-cells were found in humans identifiable by the horizontal bipolar orientation of both dendrites and axons (Figure 1E-H). The majority were located directly underneath the pial surface (Figure 1F, G) or intermingled with a heterogeneous population of GABAergic interneurons (Figure 1E).

Revised version: Remarkably, and in contrast to rodents, a persistent subpopulation of excitatory CR-cells was found in humans, identifiable by the horizontal bipolar orientation with intralaminar long-range horizontal axonal collaterals with a transcolumnar projection (Figure 1E-H; Retzius 1893, 1894; Anstötz et al. 2014; reviewed by Martinez-Cerdeno and Noctor 2014). The majority were located directly underneath the pial surface (Figure 1F, G) or intermingled with a heterogeneous population of GABAergic interneurons (Figure 1E).

Earlier version: In the human TLN, GABAergic interneurons constitute the predominant class of neurons in L1, exhibiting significant structural and functional heterogeneity (Figure 1D, E; see also Jiang et al, 2014, Verhoog et al. 2016; Boldog et al. 2018; Obermayer et al. 2018; Kwon et al. 2019; Schuman et al. 2019, reviewed by Huang et al. 2024) which were found throughout both sublaminae and were sometimes organized in clusters of 3-6 neurons (Figure 1D). Beside CR-cells and GABAergic interneurons, also degenerating neurons, identifiable by their shrunken and dark appearance (Figure 1A, I) were found in both sublaminae although their numbers varied substantially between brain tissue samples.

Revised version: In the human TLN, GABAergic interneurons constitute the predominant class of neurons in L1, exhibiting significant structural and functional heterogeneity (Figure 1D, E; see also Jiang et al, 2013, Verhoog et al. 2016; Boldog et al. 2018; Obermayer et al. 2018; Kwon et al. 2019; Schuman et al. 2019, reviewed by Huang et al. 2024) which were found throughout both sublaminae and were sometimes organized in clusters of 3-6 neurons (Figure 1D). Beside CR-cells and GABAergic interneurons, also degenerating neurons, identifiable by their shrunken and dark appearance (Figure 1A, I) were found in both sublaminae although their numbers varied substantially between brain tissue samples. L1b is primarily dominated by dendritic and synaptic profiles due to a massive increase in synaptic complexes (Figures 1H, 2B). Furthermore, L1b is infiltrated by apical dendrites of L2, L3 and L5 pyramidal neurons interspersed with smaller dendritic segments representing apical oblique dendrites (Figures 1I, 2B). The majority of synaptic complexes in L1 were axo-spinous (∼80%, Figure 2 B, D, E), the remainder were axo-dendritic, some of which but not all were regarded as putative GABAergic terminals (Figure 2C; but see Silver et al. 2003).

Removed in the revised version: L1b is primarily dominated by dendritic and synaptic profiles. due to a massive increase in synaptic complexes (Figures 1H, 2B). Furthermore, L1b is infiltrated by apical dendrites of L2, L3 and L5 pyramidal neurons which are interspersed with smaller dendritic segments representing apical oblique dendrites (Figures 1I, 2B). The majority of synaptic complexes in L1 were axo-spinous (∼80%, Figure 2 B, D, E), the remainder were axo-dendritic some of which but not all were regarded as putative GABAergic terminals (Figure 2C; but see Silver et al. 2003).

L1b can be separated from L2 by the sudden increase of neuronal cell bodies and thick apical dendrites originating at the upper pole of the somata of pyramidal neurons (Figures 1A, 2B). At the L1b and L2 transition more and more apical and apical oblique dendrites of different caliber were observable.

Synaptic organization of L1 in the human TLN

The primary objective of this study was to investigate the synaptic organization of excitatory SBs in L1 in the human TLN, focusing on synaptic parameters that represent possible morphological correlates of synaptic transmission, efficacy, strength and plasticity. The quantitative analysis was separated for both sublaminae to look for possible sublaminar-specific differences. Beside the overall geometry of SBs and the number and size of mitochondria, in particular the size of the PreAZ and PSD constituting the AZ, the morphological equivalent to a functional neurotransmitter release site, and the organization of the three functionally defined pools of SVs, namely the readily releasable (RRP), the recycling (RP) and resting pool, were quantified.

To achieve this, a total of 361 SBs was completely reconstructed out of four patients in six series containing 100-150 ultrathin sections/series. A total of 190 SBs and 184 AZs (L1a) and 171 SBs and 165 AZs (L1b) using non-affected neocortical access tissue (see Material and Methods; Source Data 02) were completely 3D-reconstructed and quantified using fine-scale TEM, digital TEM images and subsequent 3D-volume reconstructions.

Excitatory and inhibitory synaptic complexes in L1a und L1b were formed by either presynaptic en passant or endterminal boutons with their prospective postsynaptic target structures dendritic shafts (Figures 2B C, D, 3A-C) or spines of different caliber and type (Figures 2B, D, E, 3A-C). In both sublaminae, SBs were predominantly found on dendritic spines (∼81%; L1a) and ∼79% (L1b), the remainder were located on dendritic shafts. The majority of SBs on spines were located on mushroom spines (L1a: ∼83%, L1b: ∼73%), a smaller fraction on stubby spines (L1a: ∼11%, L1b: ∼9%) or on thin elongated spines (L1a: 6%, L1b: ∼22%; Figure 2E), the remainder were not classifiable.

Removed in the revised version: Since our focus was on excitatory SBs, inhibitory terminals were not further analyzed in detail. Inhibitory SBs were distinguished from excitatory SBs by their smaller SVs which also has a more spherical appearance and the smaller sized PreAZ and the lack of a prominent PSD (Figure 2D). Infrequently, GABAergic and glutamatergic terminals were found on the same dendrite (Figure 2C, D) or spine (see for example Wittner et al. 2001) but rare by ∼2-3%. (Kwon et al. 2019). In turn, also not all SBs located on dendritic shafts were GABAergic (see for example Silver et al. 2003).

Removed in the revised version: The ratio between excitatory vs. inhibitory SBs was between 10-15% although with layer-specific differences. Our findings are in good agreement with recent publications by Shapson-Coe et al. (2024) and by Cano-Astorga et al. (2023, 2024) which also estimated the number of inhibitory SBs in different layers of the temporal lobe neocortex by10-15%. However, no significant differences were found in the content of inhibitory SBs between the two sublaminae L1a and L1b.

Removed in the revised version: Some SBs in both sublaminae were seen to establish either two or three SBs on the same spine, spines of other origin or dendritic shafts. Remarkably, ∼90% of spines in L1a and L1b contained a spine apparatus (Figure 2E), a specialized form of the endoplasmic reticulum, structures that increase spine motility and may also stabilize the synaptic complex during signal transduction (Deller et al. 2003; reviewed by Knott and Holtmaat 2008).

Removed in the revised version: Furthermore, in numerous SBs, so-called dense-core vesicles (DCVs) were distributed throughout the presynaptic terminal intermingled with the population of SVs (Figure 3B, C, E).

3D-volume reconstructions of SBs and vesicle pools in L1a and L1b of the human TLN.

Earlier version: A-D, Here, all SBs are colored in transparent yellow. Subelements: PreAZs (red), SVs (green), DCVs (magenta), mitochondria (white). The postsynaptic target structures are given in blue A, Six SBs (sb1-sb6) terminating at different locations on a dendritic segment in L1b, two of which (syn3 and syn4) were located on spines, the remaining on the dendritic shaft. Scale bar 1 μm. B, Large SB terminating on two dendritic segments of different shape and size (de1, de2) in L1a. C, SB synapsing on a dendrite (de) and spine (sp) in L1b containing several mitochondria associated with the pool of SVs. D, E, 3D-volume reconstructions of the total pools of SVs (green dots). Note the different shape and size of the PreAZs (red) with either a non-perforated macular (D) or perforated (E) appearance and the comparably large total pool size. Note the large DCVs intermingled with the pool of SVs.

Revised version: A-C, Here, all SBs are outlined by a yellow contour. Subelements: PreAZs (red), SVs (green), DCVs (magenta), mitochondria (white). The postsynaptic target structures are given in blue A, Six SBs (sb1-sb6) terminating at different locations on a dendritic segment in L1b, two of which (syn3 and syn4) were located on spines, the remaining on the dendritic shaft. Scale bar 1 μm. B, Large SB terminating on two dendritic segments of different shape and size (de1, de2) in L1a. C, SB synapsing on a dendrite (de) and spine (sp) in L1b containing several mitochondria associated with the pool of SVs. D, E, 3D-volume reconstructions of the total pools of SVs (green dots). Note the different shape and size of the PreAZs (red) with either a non-perforated macular (D) or perforated (E) appearance and the comparably large total pool size. Note the large DCVs intermingled with the pool of SVs.

Added in the revised version: Since our focus was on excitatory SBs, inhibitory terminals were not further analyzed in detail. Inhibitory SBs were distinguished from excitatory SBs by their smaller SVs, which have a more spherical appearance, a smaller sized PreAZ, and the lack of a prominent PSD (Figure 2D). Infrequently, GABAergic and glutamatergic terminals were found on the same dendrite (Figure 2C, D) or spine (see for example Wittner et al. 2001) but less frequently at ∼2-3%. (Kwon et al. 2019). In turn, not all SBs located on dendritic shafts were GABAergic (see for example Silver et al. 2003).

Added in the revised version: The ratio between excitatory vs. inhibitory SBs was between 10-15% although with layer-specific differences. Our findings are in good agreement with recent publications by Shapson-Coe et al. (2024) and by Cano-Astorga et al. (2023, 2024) in which the number of inhibitory SBs in different layers of the temporal lobe neocortex was estimated to be10-15%. No significant differences were found in the content of inhibitory SBs between the two sublaminae L1a and L1b.

Added in the revised version: Some SBs in both sublaminae were seen to establish either two or three synapses on the same spine, spines of other origin, or dendritic shafts. Remarkably, ∼90% of spines in L1a and L1b contained a spine apparatus (Figure 2E), a specialized form of the endoplasmic reticulum, structures that increase spine motility and may also stabilize the synaptic complex during signal transduction (Deller et al. 2003; reviewed by Knott and Holtmaat 2008).

Remarkably, beside synaptic complexes, also ‘special connections’ in L1 of the human TLN were observed. This encompassed so-called dendro-dendritic synapses and fine astrocytic processes that receive synaptic input from SBs and in turn, provide synaptic output to either dendritic shafts or spines (not shown, but see Yakoubi et al. 2019a, b; Schmuhl-Giesen et al. 2022).

Synaptic density in L1a and L1b of the human TLN

Synaptic density measurements serve as a valuable tool for quantitatively characterizing the synaptic organization of a specific cortical layer, assessing the connectivity rate, and identifying potential inter-individual differences among patients within the human TLN.

The average density of synaptic complexes in L1 was comparatively high with a value of 5.26*108 ± 8.44*107, 5.52*108 ± 1.25*108 synapses/mm3 in L1a and 5.01*108 ± 1.40*107 synapses/mm3 in L1b, ranging from 4.08*108 to 6.31*108 (L1a) and from 4.87*108 to 5.15*108 (L1b), respectively (Supplemental Table 1; Source Data 01). Strikingly, a huge inter-individual variability was found, although with no significant difference between L1a and L1b.

Therefore, L1 exhibited a comparably high density of synaptic complexes, indicating a relatively robust and thus strong connectivity of neurons within L1 of the human TLN.

Geometry and size of SBs and mitochondria in the human TLN

Overall, L1 SBs were on average medium-sized, with a mean surface area of 5.48 ± 1.40 µm2, and a mean volume of 0.50 ± 0.19 µm3, with a slight difference in size between L1a and L1b. Interestingly, the variability in both surface area and volume of SBs was comparably small in L1a and L1b as indicated by a low CV and variance (Table 1; Source Data 02) regardless of their target structures. SBs in L1 were comparable in size with SBs in L5, L6, but ∼2-fold larger than those in L4 of the human TLN (p ≤ 0.001; see also Yakoubi et al. 2019a, b: Schmuhl-Giesen et al. 2022).

Mitochondria play a pivotal role in synaptic transmission and plasticity (reviewed by Dallerac et al. 2018; see also Discussion). In L1 SBs either no (Figures 2B, 3A, E) or several mitochondria (Figures 2, 3C, D; range 0 to 8) of different shape and size were observed. When present mitochondria had a volume of 0.04 ± 0.02 µm3 in L1a and 0.05 ± 0.01 µm3 in L1b, respectively (Table 1; Source Data 02). Mitochondria contributed with ∼7% (L1) to the total bouton volume of SBs, similar to values in L6 (∼6%; Schmuhl-Giesen et al. 2022), but with a ∼1.5-fold lower percentage as those in L5 (∼12%; Yakoubi et al. 2019a) and L4 (∼13%; Yakoubi et al. 2019b).

A good correlation between the volume of the SBs and that of mitochondria was found for L1a (R2: 0.7528; Figure 4B) and for L1b (R2: 0.5992; Figure 4B). A weak correlation was also found for L6 (Schmuhl-Giesen et al. 2022) but a strong one in both L4 (Yakoubi et al. 2019a) and L5 (Yakoubi et al. 2019b) suggesting a layer-specific difference in the content of mitochondria in individual SBs.

Correlation plots for structural and synaptic parameters characterizing L1 SBs.

Correlation plots showing the strength of correlations between structural and synaptic parameters (Source Data 02). Correlations dots and regression lines for L1a are given in black and that for L1b in red. A, Surface area of SBs vs. surface area of PreAZs; B, Volume of SBs vs. volume of mitochondria; C, Surface area of PreAZs vs. total pool of SVs; D, Surface area of PreAZs vs. p10 nm RRP; E, Surface area of PreAZs vs. p20 nm RRP; F, Surface area of PreAZs vs. p60-p200 nm RP; G, Surface area of PreAZs vs. p200 nm resting pool; H, Surface area of PreAZs vs. p500 nm resting pool.

Structural composition of AZs in L1a and L1b excitatory SBs in the human TLN

The number, size, and shape of the AZ, composed of the PreAZ and PSD, is one key structural determinant in synaptic transmission and plasticity (Südhof 2002; Matz et al. 2010; Holderith et al. 2012) as it represents the docking area of SVs (PreAZs) and the receiving area of neurotransmitter quanta at the PSD (Südhof 2012).

Earlier version: The majority (∼98%) of SBs in L1a and L1b had only a single (Figures 2C-E, 3A-C, E) at most two or three AZs (Figure 3D). Beside very large, spanning the entire pre- and postsynaptic apposition zone, also smaller AZs covering only a fraction of the pre- and postsynaptic apposition zone were found. The majority of AZs (∼75%) were of the macular, non-perforated type (Figures 2C-E, 3A-D), the remainder showed either a perforation in the PreAZ, PSD or both (Figure 3E).

Revised version: The majority (∼98%) of SBs in L1a and L1b had only a single (Figures 2C-E, 3A-C, E) at most two or three AZs (Figure 3D). Beside very large, spanning the entire pre- and postsynaptic apposition zone, also smaller AZs covering only a fraction of the pre- and postsynaptic apposition zone were found.

Earlier version: On average, PreAZs were 0.18 ± 0.08 µm2 in surface area in L1a and 0.22 ± 0.04 µm2 in L1b with only a slight variability between both sublaminae. The surface area of PSDs was quite similar to that of PreAZs 0.21 ± 0.01 µm2 and was in L1a and 0.22 ± 0.05 µm2 in L1b, respectively (Table 1; Source Data 02, sheet 01). Thus, the similar size of the PreAZ and PSD showed a nearly perfect overlap at the pre- and postsynaptic apposition zone. L1 AZs did not show a large variability in size as indicated by the low SD, CV, and variance (Table 1).

Revised version: On average, PreAZs were 0.18 ± 0.08 µm2 in surface area in L1a and 0.22 ± 0.04 µm2 in L1b with only a slight variability between both sublaminae. The surface area of PSDs was quite similar to that of PreAZs 0.21 ± 0.01 µm2 and was in L1a and 0.22 ± 0.05 µm2 in L1b, respectively (Table 1; Source Data 02, sheet 01). Thus, the similar size of the PreAZ and PSD showed a nearly perfect overlap at the pre- and postsynaptic apposition zone. L1 AZs did not show a large variability in size as indicated by the low SD, CV, and variance (Table 1). Remarkably, AZs exhibited a wide range of shapes in addition to different sizes. Two types of AZs were found: The majority in L1 (∼80%) were of the macular and non-perforated type (Figure 3D), the remainder being either horseshoe-shaped or ring-like and perforated (Figure 3E).

Earlier version: Notably, our analysis revealed only a weak correlation between the surface area of SBs and that of PreAZs for L1a and L1b SBs (Figure 4A). These findings imply that the size of SBs and PreAZs is independently regulated from each other.

Revised version: Notably, our analysis revealed only a weak correlation between the surface area of SBs and that of PreAZs for L1a and L1b SBs (Figure 4A). These findings imply that the size of SBs and PreAZs is regulated independently.

Earlier version: Measurements of the size of the synaptic cleft were performed on random EM images taken from the series using OpenCAR or online directly at the TEM using ImageSP (Fa. Tröndle, Moorenweis, Germany). Only synaptic clefts cut perpendicular to the PreAZ and PSD were included in the sample. The distance between the outer edge of the PreAZ and PSD was measured at the two lateral edges and separately at the central region of the synaptic cleft; the two values of the lateral edges were averaged for each cleft measurement. Finally, a mean ± SD was calculated for both the lateral and central region over all synaptic clefts analyzed per patient.

Revised version: The size of the synaptic cleft (Table 1; Source Data 02, sheet 03) is an important measure at AZs since its size (diameter) critically determines the temporal and spatial neurotransmitter concentration. Measurements of the size of the synaptic cleft were performed on random EM images taken from the series using OpenCAR or online directly at the TEM using ImageSP (Fa. Tröndle, Moorenweis, Germany). Only synaptic clefts cut perpendicular to the PreAZ and PSD were included in the sample. The distance between the outer edge of the PreAZ and PSD was measured at the two lateral edges and separately at the central region of the synaptic cleft; the two values of the lateral edges were averaged for each cleft measurement. Finally, a mean ± SD was calculated for both the lateral and central region over all synaptic clefts analyzed per patient.

Earlier version: The size of the synaptic cleft (Table 1; Source Data 02, sheet 03) is an important measure at AZs since its size (diameter) critically determines the temporal and spatial neurotransmitter concentration. The synaptic cleft size were slight, but non-significantly different between the two sublaminae and similar to findings in L4, L5 and L6 (Yakoubi et al. 2019a, b: Schmuhl-Giesen et al. 2022).

Revised version: The size of the synaptic cleft did not differ significantly between the two sublaminae and was similar to the results in L4, L5 (Yakoubi et al. 2019a, b) and L6.Schmuhl-Giesen et al. 2022).

Organization of the pools of SVs in L1a and L1b excitatory SBs of the human TLN

SVs are the other key structure in neurotransmitter storage and release, hence they play a fundamental role in synaptic transmission and in the modulation of short- and long-term synaptic plasticity (Südhof 2002, 2012). Three different pools of SVs are functionally defined: the RRP, the RP, and the resting pool. Synaptic efficacy, strength, mode, and probability of release (Pr) are regulated by these pools (Schikorski and Stevens 2001; Silver et al. 2003; Rizzoli and Betz 2004; Saviane and Silver 2006; Schikorski 2014; Watanabe et al. 2014; Neher 2015; Vaden et al. 2019; reviewed by Rizzoli and Betz 2005; Denker and Rizzoli 2010; Chamberland and Toth 2016).

Earlier version: In general, SVs were distributed throughout the entire terminal in ∼95% of the population of SBs investigated (Figure 2, 3), the remainder showed a more cluster-like distribution. Two different types of vesicles were found: (1) Small clear SVs with an average diameter of 25.03 ± 5.27 nm (Table 1; Source Data 02, sheet 01) and (2) DCVs (Figure 3B, C, E; Source Data 02, sheet 01) with an average diameter of 57.97 ± 6.45 nm with no significant difference between both sublaminae. DCVs were intermingled with the population of SVs throughout the entire SB (Figure 3B, C, E).

Revised version: In general, SVs were distributed throughout the entire terminal in ∼95% of the population of SBs investigated (Figure 2, 3), the remainder showed a more cluster-like distribution. Two different types of vesicles were found: (1) Small clear SVs with an average diameter of 25.03 ± 5.27 nm ranging up to ∼40 nm (Table 1; Source Data 02, sheet 01). Using TEM tomography, the SV diameter was on average 30.29 ± 4.67 nm ranging up to ∼47 nm (n = 38 SBs with a total of 5.834 SVs; Table 1; Source Data 03, sheet 05). (2) DCVs (Dense Core Vesicles; Figure 3B, C, E; Source Data 02, sheet 01) with an average diameter of 57.97 ± 6.45 nm with no significant difference between both sublaminae. DCVs were intermingled with the population of SVs throughout the entire SB (Figure 3B, C, E). DCVs play an important role in endo- and exocytosis by releasing various co-transmitters, such as neuropeptides, ATP, noradrenalin, and dynorphin (Ghijsen and Leenders 2005) or in synaptogenesis (Sorra et al. 2006) and the build-up of PreAZs by releasing Piccolo and Bassoon (Zhai et al. 2001; Shapira et al. 2003) In addition, they are involved in clustering SVs at the PreAZs (Mukherjee et al. 2010; Watanabe et al. 2014).

Removed in the revised version: DCVs play an important role in endo- and exocytosis, the build-up of PreAZs by releasing Piccolo and Bassoon (Schoch and Gundelfinger 2006; Murkherjee et al. 2010), or in clustering SVs at the PreAZs (Watanabe et al. 2014, Murkherjee et al. 2010). In addition, various co-transmitters, such as neuropeptides, ATP, noradrenalin, and dynorphin were identified in large DCVs (Ghijsen and Leenders 2005).

Earlier version: The distribution pattern of SVs made it impossible to morphologically distinguish between the three functionally defined pools of SVs, except for the so-called ‘docked’ and fused vesicles primed to the PreAZ (Rizzoli and Betz, 2004; Neher 2015; reviewed by Rizzoli and Betz 2005; Chamberland and Toth 2015).

Revised version: The distribution pattern of SVs made it impossible to morphologically distinguish between the three functionally defined pools of SVs, with the exception of the so-called ‘docked’ ones, that are fused to the PreAZ, and the omega-shaped bodies, that have already opened and released quanta of neurotransmitter (Rizzoli and Betz, 2004; Denker and Rizzoli 2010; Neher 2015; reviewed by Rizzoli and Betz 2005; Chamberland and Toth 2015). However, an attempt was made to define a morphological correlate for the three functionally defined pools.

Earlier version: Thus, a perimeter analysis (a modified Sholl-analysis) was used to determine the exact location of each SV from the PreAZ in individual terminals (for criteria to define SV pools see Rizzoli and Betz 2005). The RRP was defined as a distance (perimeter p) from ≤10 nm and ≤20 nm from the PreAZ. This criteria for the RRP was chosen because both values are less than an SV diameter and hence may represent the RRP from which SVs could be easily and rapidly recruited. The RP maintains release during moderate stimulation and was estimated close (p60-p200 nm) to the PreAZ. All SVs further than 200 nm from the PreAZ were regarded to constitute the resting pool which acts as a deposit of SVs only used by intense and/or repetitive stimulation.

Revised version: For this purpose, a perimeter analysis (a modified Sholl-analysis) of each SV in individual SBs was carried out. The perimeter was defined as the minimum distance between each SV membrane and the PreAZ membrane (Sätzler et al. 2002; for criteria to define SV pools see also Rizzoli and Betz 2005; Denker and Rizzoli 2010). The putative RRP was defined as a perimeter (p) p≤10 nm and p≤20 nm to the PreAZ. These criteria for the RRP from which SVs could be easily and rapidly recruited were chosen because both values are less than a SV diameter apart from the PreAZ. The putative RP maintains release during moderate stimulation and was classified at a distance from the PreAZ between p60-p200 nm. All SVs that were further than 200 nm from the PreAZ were regarded to constitute the putative resting pool which acts as a deposit of SVs only used by intense and/or repetitive stimulation.

Earlier version: On average, the total pool of SVs was 2958.62 ± 1940.51 in L1a vs. 3903.32 ± 1852.23 in L1b with an average of 3430.97 ± 1773.77 for L1 (Table 1; Source Data 02). Remarkably, the total pool in L1b was significantly larger (p ≤ 0.01) by ∼1.3-fold when compared to L1a. A huge variability in total pool size exists in both sublaminae (Figure 3D, E) ranging from 758 to 6542 in L1a and 1507 to 15060 in L1b as indicated by the SD, IQR and variance (Table 1). SVs contributed with ∼5% (0.03 µm3) in L1a and ∼4% (0.02 µm3) in L1b to the total volume of SBs. However, no significant difference between both sublaminae was observed (Table 1; Source Data 02).

Revised version: On average, the total pool of SVs was 2958.62 ± 1940.51 in L1a vs. 3903.32 ± 1852.23 in L1b with an average of 3430.97 ± 1773.77 for L1 (Table 1; Source Data 02). Remarkably, the total pool in L1b was significantly larger (p≤0.01) by ∼1.3-fold when compared to L1a. In both sublaminae (Figure 3D, E), there is a huge variability in total pool size, ranging from 758 to 6542 in L1a and 1507 to 15060 in L1b, as indicated by the SD, IQR and variance (Table 1). SVs contributed with ∼5% (0.03 µm3) in L1a and ∼4% (0.02 µm3) in L1b to the total volume of SBs. However, no significant difference between both sublaminae was observed (Table 1; Source Data 02). Interestingly, no correlation between the volume of SBs vs. the total pool size was found for L1a (R2: 0.2490), but a good correlation for L1b (R2: 0.5700). Furthermore, no correlation was found between the PreAZ vs. the total pool size (Figure 4C; Source Data 02).

Removed in the revised version: Interestingly, no correlation between the volume of SBs vs. the total pool size was found for L1a (R2: 0.2490), but a good correlation for L1b (R2: 0.5700). Furthermore, no correlation was found between the PreAZ vs. the total pool size (Figure 4C; Source Data 02).

Earlier version: The average number of SVs in the putative RRP/PreAZ for L1was 3.60 ± 4.24 SVs at p10 nm (L1a: 5.90 ± 5.05 SVs and L1b: 1.29 ± 1.87 SVs) with a significant difference (p ≤ 0.01) by ∼4.5-fold between the two sublaminae (Table 1; Source Data 02). The putative RRP/PreAZ at the p20 nm criterion was on average 19.05 ±17.23 SVs (L1a: 25.04 ± 21.09 SVs and L1b: 13.07 ± 13.87 SVs) and thus nearly 2-fold larger for L1a. Interestingly, both RRPs in L1a and L1b were characterized by a comparably low variability as indicated by the SD, CV, and variance (Table 1) that may point to sublaminae-specific differences in Pr, synaptic efficacy, strength and paired-pulse behavior at individual SBs.

Revised version: The average number of SVs in the putative RRP/PreAZ for L1was 3.60 ± 4.24 SVs at p10 nm (L1a: 5.90 ± 5.05 SVs and L1b: 1.29 ± 1.87 SVs) with a significant difference (p ≤ 0.01) by ∼4.5-fold between the two sublaminae (Table 1; Source Data 02). The putative RRP/PreAZ at the p20 nm criterion was on average 19.05 ±17.23 SVs (L1a: 25.04 ± 21.09 SVs and L1b: 13.07 ± 13.87 SVs) and thus nearly 2-fold larger for L1a, although with a huge variability, but even though may point to sublaminae-specific differences in Pr, synaptic efficacy, strength and paired-pulse behavior at individual SBs.

Earlier version: No correlation between the RRP at both the p10 nm and p20 nm criterion and the surface area of PreAZs was found for both sublaminae. (Figure 4D, E).

Revised version: No correlation between the putative RRP at both the p10 nm and p20 nm criterion and the surface area of PreAZs was found for both sublaminae. (Figure 4D, E).

Earlier version: The RP/PreAZ at the 60-200 nm perimeter criterion was 390.12 ± 286.89 SVs in L1a and 535.88 ± 321.16 in L1b with no significant difference for both sublaminae (Table 1; Source Data 02, sheet 01). It should be noted however, that the variability of the RP is relatively large in both sublaminae as indicated by the SD, median, IQR and variance (Table 1).

Revised version: The putative RP/PreAZ at the 60-200 nm perimeter criterion was 390.12 ± 286.89 SVs in L1a and 535.88 ± 321.16 in L1b with no significant difference for both sublaminae (Table 1; Source Data 02, sheet 01). It should be noted however, that the variability of the putative RP is relatively large in both sublaminae as indicated by the SD, median, IQR and variance (Table 1).

Earlier version: Again, no correlation was observed between the SVs in the p60-200 nm RP and the surface area of the PreAZs (L1a, Figure 4F) and for L1b (Figure 4F; Source Data 02, sheet 01), respectively.

Revised version: Again, no correlation was observed between the SVs in the putative p60-200 nm RP and the surface area of the PreAZs (L1a, Figure 4F) and for L1b (Figure 4F; Source Data 02, sheet 01), respectively.

Earlier version: The resting pool of SVs contained on average 2515.02 ± 1588.99 SVs in L1a and 3277.97 ± 1480.96 SVs in L1b with a significant difference between the two sublaminae (p ≤ 0.01, Table 1; Source Data 02, sheet 01). The number of SVs in the resting pool in L1b was about 1.3-fold larger when compared to L1a. Again, also the resting pool showed a large variability for both sublaminae as indicated by the SD, CV, IQR and variance (Table 1).

Revised version: The putative resting pool of SVs contained on average 2515.02 ± 1588.99 SVs in L1a and 3277.97 ± 1480.96 SVs in L1b with a significant difference between the two sublaminae (p ≤ 0.01, Table 1; Source Data 02, sheet 01). The number of SVs in the putative resting pool in L1b was about 1.3-fold larger when compared to L1a. Again, also the putative resting pool showed a large variability for both sublaminae as indicated by the SD, CV, IQR and variance (Table 1).

Earlier version: No correlation was observed between the SVs in the p200 nm and p500 nm resting pool and the surface area of the PreAZs in L1a and L1b (Figure 4G, H; Source Data 02, sheet 01).

Revised version: No correlation was observed between the SVs in the putative p200 nm and p500 nm resting pool and the surface area of the PreAZs in L1a and L1b (Figure 4G, H; Source Data 02, sheet 01).

Earlier version: Neither for the total nor for the three different pools of synaptic vesicles significant differences between SBs terminating on dendritic spines or shafts were found (Source Data 02, sheet 04).

Revised version: Neither for the total nor for the three different putative pools of synaptic vesicles significant differences between SBs terminating on dendritic spines or shafts were found (Source Data 02, sheet 04).

TEM tomography of L1 excitatory SBs in the human TLN

It is still controversially discussed whether so-called ‘docked’ SVs or omega-shaped bodies represent the RRP. To compare our results from the perimeter analysis for the p10 nm RRP with that of ‘docked’ SVs, high-resolution TEM tomography was carried out. Only SBs where the AZ could be followed from its beginning to its end in individual tilt-series and where the AZ was cut perpendicular through the PreAZ, PSD and the synaptic cleft were analyzed. In L1a (50 SBs, on dendritic shafts = 25 SBs; on spines = 25 SBs) and L1b (50 SBs, on dendritic shafts = 25 SBs; on spines = 25 SBs) the number of ‘docked’ SVs was analyzed at the PreAZ (Figure 5; Table 2; Source Data 03).

EM tomography of SBs in L1a and L1b in the human TLN.

Earlier version: A, Large SB (sb) terminating with a single AZ (arrowheads) on a terminal tuft dendritic shaft (de) in L1a. B, SB (sb) synapsing on a stubby spine (sp) with a prominent spine apparatus (framed area) and a thick dendritic segment (de) in L1b. Scale bars 1 µm. C, Large mushroom (sp1) and a small spine head (sp2) receiving input from three SBs (sb1, sb2, sb3) and a single SB (sb4) in L1b. Note the large AZs (arrowheads) and the prominent spine apparatus (red framed area) in the mushroom spine. Scale bar 0.5 µm. D, E, F, High-power images of ‘docked’ SVs (arrowheads) taken from a tilt-series through an individual PreAZ at a L1a spine SB. Note the so-called omega-shaped bodies in (E, F) pointing to the already opening and release of glutamate quanta. Scale bars 0.1 µm.

Revised version: A, Large SB (sb) terminating with a single AZ (arrowheads) on a terminal tuft dendritic shaft (de) in L1a. B, Two adjacent SBs synapsing either (sb1) on a thick dendritic segment (de) or (sb2) on a stubby spine (sp) with a prominent spine apparatus (framed area) in L1b. Scale bars 1 µm. C, Large mushroom (sp1) and a small spine head (sp2) receiving input from three SBs (sb1, sb2, sb3) and a single SB (sb4) in L1b. Note the large AZs (arrowheads) and the prominent spine apparatus (red framed area) in the mushroom spine. Scale bar 0.5 µm. D, E, F, High-power images of ‘docked’ SVs (arrowheads) taken from a tilt-series through an individual PreAZ at a L1a spine SB. Note the so-called omega-shaped bodies in (E, F) pointing to the already opening and release of glutamate quanta. Scale bars 0.1 µm.

Earlier version: The results for L1a and L1b were two-fold: First, and in agreement with already published results (Yakoubi et al. 2019a, b; Schmuhl-Giesen et al. 2022) only in a minority (<1%) of all PreAZs analyzed, regardless of their target structures, a dendritic shaft (Figure 5A) or spine (Figure 5B, C), no ‘docked’ SVs were observed. Second, the majority of PreAZs (∼98%) contained more than 2, the most 8 (L1a) and 6 (L1b) ‘docked’ SVs. This finding strongly supports multivesicular release of ‘docked’ SVs in L1a and L1b SBs in line with findings in L4, L5 and L6 SBs (Yakoubi et al. 2019a, b; Schmuhl-Giesen et al. 2022; see also Figure 5D-F).

Revised version: The results for L1a and L1b were two-fold: First, in agreement with already published results (Yakoubi et al. 2019a, b; Schmuhl-Giesen et al. 2022) only in a minority (<1%) of all PreAZs analyzed, regardless of their target structures, a dendritic shaft (Figure 5A) or spine (Figure 5B, C), no ‘docked’ SVs were observed, which is in contrast to our perimeter analyses showing several boutons without vesicles within the p10 nm criterion. Second, the majority of PreAZs (∼98%) contained more than 2, the most 8 (L1a) and 6 (L1b) ‘docked’ SVs. This finding strongly supports multivesicular release of ‘docked’ SVs in L1a and L1b SBs in line with findings in L4, L5 and L6 SBs (Yakoubi et al. 2019a, b; Schmuhl-Giesen et al. 2022; see also Figure 5D-F).

Earlier version: On average, 3.71 ± 1.38 ‘docked’ SVs (L1a) and 3.42 ± 1.34 ‘docked SVs’ (L1b) were found at individual PreAZs with similar values for both sublaminae. Furthermore, no significant difference was found for ‘docked’ SVs between shaft vs. spine synapses in L1. Strikingly, the number of ‘docked’ SVs at L1a PreAZs was ∼2-fold smaller but nearly 3-fold larger in L1b when compared to the results of our quantitative perimeter analysis for the p10 nm criterion and were also significantly different (p < 0.001).

Revised version: On average, 3.71 ± 1.38 ‘docked’ SVs (L1a) and 3.42 ± 1.34 ‘docked SVs’ (L1b) were found at individual PreAZs with similar values for both sublaminae. Furthermore, no significant difference was found for ‘docked’ SVs between shaft vs. spine synapses in L1. The number of ‘docked’ SVs at L1a PreAZs was ∼2-fold smaller but nearly 3-fold larger in L1b when compared to the results of our quantitative perimeter analysis for the putative RRP using the p10 nm criterion. In contrast to the p10 nm criterion with high inter- and intra-individual variabilities (Suppl. Fig. 2), the number of ‘docked’ SVs showed a relatively low variability as indicated by the SD, CV, variance, and skewness. They were different (p ≤ 0.001) from the values estimated with the p10 nm perimeter analysis.

Earlier version: The number of ‘docked’ SVs showed a relatively low variability as indicated by the SD, CV, variance, and skewness, but were different from the values estimated with the p10 nm perimeter analysis. There seems to be a tendency that larger PreAZs contained more ‘docked’ SVs providing a larger ‘docking’ area allowing the recruitment of more SVs.

Revised version: There seems to be a tendency that larger PreAZs contained more ‘docked’ SVs providing a larger ‘docking’ area allowing the recruitment of more SVs.

Earlier version: In summary, a notable disparity difference was observed between values obtained for L1 using the p10 nm criterion of the perimeter- and the EM tomography analysis. Additionally, the RRP in L1 and L6 was ∼2-4-times smaller when compared to values in L4 and L5 (Yakoubi et al. 2019a, b; Schmuhl-Giesen et al. 2022) pointing towards a layer-specific regulation of the RRP. Moreover, the results suggest that not only ‘docked’ SVs but also those very close to the PreAZ should be considered to belong to the RRP.

Revised version: In summary, a notable disparity difference was observed between values obtained for L1 using the p10 nm criterion of the perimeter- and the TEM tomography analysis. Additionally, the putative RRP in L1 (this study) and L6 (Schmuhl-Giesen et al. 2022) was ∼2-4-times smaller when compared to values in L4 (Yakoubi et al. 2019b) and L5 (Yakoubi et al. 2019a) pointing towards a layer-specific regulation of the putative RRP. Moreover, the results suggest that not only ‘docked’ SVs but also those very close to the PreAZ should be considered to belong to the putative RRP.

Astrocytic coverage of L1 SBs in the human TLN

Earlier version: Astrocytes, by directly interacting with synaptic complexes thus forming the ‘tripartite’ synapse play a pivotal role in the induction, maintenance and termination of synaptic transmission by controlling the spatial and temporal concentration of neurotransmitter quanta in the synaptic cleft (reviewed by Dellerac et al. 2018). Astrocytic profiles were identified by their irregular stellate shape, relatively clear cytoplasm, numerous glycogen granules and bundles of intermediate filaments (for criteria see Peters et al. 1991; Ventura and Harris 1999; see also Figure 6A, B). Astrocytes and their fine processes formed a relatively dense (in L4, L5), but a comparably loose network within the neuropil in L1 (Figure 6A-C; Source Data 04, sheet 01) by contributing with ∼20% to the total human TLN volume, with the highest degree in coverage in L4 but similar to findings in L6 of the human TLN (Figure 6C; Source Data 04, sheet 01).

Revised version: Astrocytes, by directly interacting with synaptic complexes thus forming the ‘tripartite’ synapse play a pivotal role in the induction, maintenance and termination of synaptic transmission by controlling the spatial and temporal concentration of neurotransmitter quanta in the synaptic cleft (reviewed by Dellerac et al. 2018). Astrocytic profiles were identified by their irregular stellate shape, relatively clear cytoplasm, numerous glycogen granules and bundles of intermediate filaments (for criteria see for example Peters et al. 1991; Ventura and Harris 1999; see also Figure 6A, B) Astrocytes and their fine processes formed a relatively dense (in L4, L5), but a comparably loose network within the neuropil in L1 (Figure 6A-C; Source Data 04, sheet 01) by contributing with ∼20% to the total human TLN volume, with the highest degree in coverage in L4 but similar to findings in L6 of the human TLN (Figure 6C; Source Data 04, sheet 01).

Astrocytic coverage of synaptic complexes in the human TLN.

Earlier version: A, EM micrograph of two adjacent SBs (sb1, sb2) terminating on two spines (sp1, sp2) in L1a. Sp2 contained a prominent spine apparatus (framed area), sb1 two DCVs (asterisks). Note that both synaptic complexes were only partially ensheathed by fine astrocytic processes (transparent green) reaching the AZs only on one side. In both synaptic complexes two ‘docked’ SVs (transparent red) at the PreAZs were found. Scale 0.25 µm. B, Stubby spine (stsp) receiving two SBs (sb1, sb2) in L1b. Here, fine astrocytic processes (transparent green) were found close to the two synaptic complexes but never reached the synaptic cleft. Note also the ‘docked’ SVs (transparent red) pointing to multivesicular release. In both images the AZs are marked by arrowheads. Scale bar 0.25 µm. C, Bar histogram showing the percentage (mean ± SD) of the volumetric fraction of astrocytic processes to the total volume in L1 to L6 (Source Data 04, sheet 01). Values for L4-L6 are taken from Schmuhl-Giesen et al. 2022.. D, Bar histogram showing the percentage (mean ± SD) of AZs covered by fine astrocytic processes either on both sides, one side or without any coverage subdivided by their location on dendritic spines or shafts (Source Data 04, sheet 02). The horizontal bars in C and D indicate significant differences *: p < 0.05; **: p < 0.01.

Revised version: A, EM micrograph of two adjacent SBs (sb1, sb2) terminating on two spines (sp1, sp2) in L1a. Sp2 contained a prominent spine apparatus (framed area), sb1 two DCVs (asterisks). Note that both synaptic complexes were only partially ensheathed by fine astrocytic processes (transparent green) reaching the AZs only on one side. In both synaptic complexes two ‘docked’ SVs (transparent red) at the PreAZs were found. Scale 0.25 µm. B, Stubby spine (stsp) receiving two SBs (sb1, sb2) in L1b. Here, fine astrocytic processes (transparent green) were found close to the two synaptic complexes but never reached the synaptic cleft. Note also the ‘docked’ SVs (transparent red) pointing to multivesicular release. In both images the AZs are marked by arrowheads. Scale bar 0.25 µm. C, Bar histogram showing the percentage (mean ± SD) of the volumetric fraction of astrocytic processes to the total volume in L1 to L6 (Source Data 04, sheet 01). Values for L4-L6 are taken from Schmuhl-Giesen et al. 2022. D, Bar histogram showing the percentage (mean ± SD) of AZs covered by fine astrocytic processes, either on both sides, on one side or without any coverage, subdivided by their location on dendritic spines or shafts in L1 and L4 (Source Data 04, sheet 02). The horizontal bars in C and D indicate significant differences *p < 0.05; **p < 0.01.

Earlier version: In general, the majority (∼70% both in spine and shaft synapses; Figure 6D, Source Data 04, sheet 02) of synaptic complexes in L1 were without astrocytic coverage close to the AZ or only partially ensheated (∼20% both in spine and shaft synapses) by fine astrocytic processes. In only a small fraction (∼10%) of both spine and shaft synapses fine astrocytic processes were observed to form the ‘tripartite’ synapse that physically isolate synaptic complexes from the surrounding neuropil. The findings in L1 are in contrast to L4 were a higher number of ‘tripartite’synaptic complexes (spines: ∼50%; shafts: ∼20%) was found. Around 25% AZs on dendritic spines and ∼ 60% on dendritic shafts were partially covered with a huge inter-individual difference. The remainder were not covered. No significant differences were found between spine and shaft synapses (Figure 6D, Source Data 04, sheet 02).

Revised version: In general, the majority (∼70% both in spine and shaft synapses; Figure 6D, Source Data 04, sheet 02) of synaptic complexes in L1 were without astrocytic coverage close to the AZ or only partially ensheated (∼20% both in spine and shaft synapses) by fine astrocytic processes. In only a small fraction (∼10%) of both spine and shaft synapses fine astrocytic processes were observed to form the ‘tripartite’ synapse that physically isolate synaptic complexes from the surrounding neuropil. The findings in L1 (L1a: n=3 patients, n=128 spine synapses, n= 48 shaft synapse; L1b: n=3 patients, n=126 spine synapses, n= 44 shaft synapses) are in contrast to L4 (n=5 patients; n=93 spine synapses, n= 41 shaft synapses) were a higher number of ‘tripartite’ synaptic complexes (spines: ∼50%; shafts: ∼20%) was found. In L4, ∼ 25% of AZs on dendritic spines and ∼ 60% on dendritic shafts were partially covered, with a huge inter-individual difference. The remainder were not covered. No significant differences were found between spine and shaft synapses in L1 or L4 (Figure 6D, Source Data 04, sheet 02).

Earlier version: Hence, it is most likely that the partial coverage or complete absence of fine astrocytic processes at the majority of synaptic complexes in L1 of the human TLN may contribute to different ‘behaviors’ of synaptic complexes in synaptic transmission and in the modulation of short-term plasticity, for example in the removal of ‘spilled’ horizontally diffusing neurotransmitter quanta favoring synaptic crosstalk (for more detail see Discussion).

Revised version: In summary, it is most likely that the partial coverage or complete absence of fine astrocytic processes at the majority of synaptic complexes in L1 of the human TLN may contribute to different ‘behaviors’ of synaptic complexes in synaptic transmission and in the modulation of short-term plasticity, for example in the removal of ‘spilled’ horizontally diffusing neurotransmitter quanta favoring synaptic crosstalk (for more detail see Discussion).

Discussion

Earlier version: This study investigated the synaptic organization of L1 in the human TLN. The quantitative 3D-models analyzed showed beside similarities significant layer-specific differences in those structural and synaptic parameters, in particular in the size of the AZ and the total pool of SVs, the RRP, RP and resting pool when compared with other already investigated layers (Yakoubi et al. 2019a, b, Schmuhl-Giesen et al. 2022). The low degree of astrocytic coverage of L1 SBs suggests that glutamate spillover and as a consequence synaptic cross talk may occur at the majority of synaptic complexes in L1

Revised version: This study investigated the synaptic organization of L1 in the human TLN. The quantitative 3D-models analyzed showed beside similarities significant layer-specific differences in structural and synaptic parameters, in particular in the size of the AZs and the total pool of SVs, the putative RRP, RP and resting pool when compared with other already investigated layers (Yakoubi et al. 2019a, b; Schmuhl-Giesen et al. 2022). The low degree of astrocytic coverage of L1 SBs suggests that glutamate spillover and as a consequence synaptic cross talk may occur at the majority of synaptic complexes in L1

Synaptic density measurements

Earlier version: Synaptic density measurements are a useful tool to describe the synaptic organization of a particular area, nuclei and even layers in different brain regions, but also the degree of connectivity underlying the computational properties of a given brain area or in a given brain network. Meanwhile numerous studies in various animal species and brain regions have performed such an analysis but data for a density of synaptic complexes in humans are still rare (but see Marco and DeFelipe 1998; Tang et al. 2001; DeFelipe et al. 1999, 2002; Alonso-Nanclares et al. 2008; Blazquez-Llorca et al. 2013; Finnema et al. 2016; Cano-Astorga et al. 2021, 2023).

Revised version: Synaptic density measurements are a useful tool to describe the synaptic organization of a particular area, nuclei and even layers in different brain regions, but also the degree of connectivity underlying the computational properties of a given brain area or in a given brain network. Meanwhile numerous studies in various animal species and brain regions have performed such an analysis but data for a density of synaptic complexes in humans are still rare (but see Marco and DeFelipe 1997; Tang et al. 2001; DeFelipe et al. 1999, 2002; Alonso-Nanclares et al. 2008; Blazquez-Llorca et al. 2013; Finnema et al. 2016; Cano-Astorga et al. 2021, 2023).

Earlier version: Strikingly, a huge layer-specific difference in the mean density of synaptic contacts was found in the human TLN (see also Supplemental Table 1). In L1, the overall density was 5.26*108 ± 8.44*107 synaptic complexes/mm3 although a great variability was observed as indicated by the SD consistent with findings in the other layers of the human TLN L4 (2.37*106 ± 2.19*106; Yakoubi et al. 2019b); L5 (3.89 *108 ± 9.12*108; Yakoubi R., PhD Thesis 2023) and L6 (4.98*107 ± 1.85*107; Schmuhl-Giesen et al. 2022) and existing data published by Alonso-Nanclares et al. 2008: 9.13 ± 0.63*108 and Tang et al. 2001: 164*1012. The density values for L1 and L5 are in the same order of magnitude but differ substantially from L4 and L6 by one or two orders of magnitude.

Revised version: Strikingly, a huge layer-specific difference in the mean density of synaptic contacts was found in the human TLN (see also Supplemental Table 1). In L1, the overall density was 5.26*108 ± 8.44*107 synaptic complexes/mm3 although a great variability was observed as indicated by the SD consistent with findings in the other layers of the human TLN L4 (2.37*106 ± 2.19*106; Yakoubi et al. 2019b); L5 (3.89 *108 ± 9.12*108; Yakoubi R., PhD Thesis 2023) and L6 (4.98*107 ± 1.85*107; Schmuhl-Giesen et al. 2022) and existing data published by Alonso-Nanclares et al. 2008: 9.13 ± 0.63*108 and Tang et al. 2001: 164*1012. The density values for L1 and L5 are in the same order of magnitude but differ substantially from L4 and L6 by one or two orders of magnitude. The huge differences in synaptic density may be explained by gender and age-specific differences and the use of different methods in estimating synaptic densities (DeFelipe et al. 1999).

Earlier version: In summary, the highest synaptic density was observed in L1, indicating a heightened level of connectivity and synaptic interaction. This suggests that L1 facilitates rapid information processing due to its robust synaptic activity.

Revised version: In summary, the highest synaptic density was observed in L1, indicating a comparably high level of connectivity and synaptic interaction. This may suggests that L1 facilitates rapid information processing due to its robust synaptic activity.

Earlier version: Important structural subelements of SBs in the human TLN Shape and size of PreAZs and PSDs

Revised version: Important structural subelements of SBs in the human TLN

Added in the revised version: Shape and size of PreAZs and PSDs

Synaptic efficacy, strength, modes of release and Pr are beside the pool of SVs, determined by the shape and size of AZs (Matz et al. 2010; Holderith et al. 2012; Südhof 2012).

Earlier version: The majority of SBs in L1 of the human TLN had a single at most three AZs that could be of the non-perforated macular or perforated type comparable with results for other layers in the human TLN (L4 Yakoubi et al. 2019b, L5: Yakoubi et al. 2019a, L6: Schmuhl-Giesen et al. 2022) but by ∼1.5-fold larger than in rodent and non-human primates (Marrone et al. 2005; Rollenhagen et al. 2015, 2018; Bopp et al. 2017; Hsu et al. 2017).

Revised version: The majority of SBs in L1 of the human TLN had a single at most three AZs that could be of the non-perforated macular or perforated type. These findings are comparable with results for other layers in the human TLN (L4: Yakoubi et al. 2019b, L5: Yakoubi et al. 2019a, L6: Schmuhl-Giesen et al. 2022) but by ∼1.5-fold larger than in rodent and non-human primate neocortex (Marrone et al. 2005; Rollenhagen et al. 2015, 2018; Bopp et al. 2017; Hsu et al. 2017).

The surface area of AZs in L1 SBs of the human TLN was on average ∼0.20 μm2. This is in good agreement with data obtained for AZs in L3 (Cano-Astorga et al. 2021), L4 (Yakoubi et al. 2019b), L5 (Yakoubi et al. 2019a), L6 (Schmuhl-Giesen et al. 2022) and in L3 (Cano-Astorga et al. 2023) of the human temporal and cingulate neocortex. However, they were ∼2- to 3-fold larger than those in mouse and non-human primates visual, motor, and somatosensory neocortex (Bopp et al. 2017; Hsu et al. 2017) but remarkably even larger than AZs in comparably large CNS terminals like the hippocampal mossy fiber bouton (Rollenhagen et al. 2007) the cerebellar mossy fiber bouton (Xu-Friedman and Regehr 2003)and the Calyx of Held in the medial nucleus of the trapezoid body (Sätzler et al. 2002).

The substantial variability in the size of AZs at individual SBs, observed both in the human TLN and in experimental animals, likely contributes to differences in synaptic efficacy, strength, Pr, quantal size, as well as the size of the RRP and RP (Südhof 2002; Matz et al. 2010, Freche et al. 2011; Holderith et al. 2012; Neher 2015; Chamerland and Toth 2016; Rollenhagen et al. 2018; Vaden et al. 2019; reviewed by Rizzoli and Betz 2005; Denker and Rizzoli 2010). However, it has to be noted that the size of the AZ is regulated as a function of activity, as shown for hippocampal SBs in the CA1 subregion (Matz et al. 2010; Holderith et al. 2012). However, the comparably large size of the AZs in the human TLN also suggest a larger ‘docking’ area allowing the fusion of more SVs and thus a larger Pr.

Added in the revised version: It has be postulated that perforated AZs contribute to synaptic efficacy (Geinismann et al. 1992; Nava et al. 2014) by an increased number of SVs at perforated synapses and plasticity (Buchs and Muller 1996; Desmond and Weinberg; 1998) and by increasing the number of AMPA- and NMDA-type glutamate receptors at perforated synapses (Ganeshina et al. 2004). However, since the number of perforated AZs in all layers of the human TLN is comparably low compared to those of the macular-non-perforated type, it is highly speculative which impact they might have here on synaptic efficacy and the modulation of synaptic plasticity.

Added in the revised version: Size of SVs obtained with TEM vs. TEM tomography

Added in the revised version: The values obtained with both TEM (25.03 ± 5.27 nm) and TEM tomography (30.29 ± 4.67 nm) in L1 are within a range published in previous studies (e.g. human TLN: Yakoubi et al. 2019a, b; Schmuhl-Giesen et al. 2022; rat hippocampal CA3 region: Rollenhagen et al. 2007; Zhao et al. 2012a; rodent hippocampal CA1 region: Harris and Sultan 1995 [lower range]; rodent somatosensory cortex: Rollenhagen et al. 2015; 2018; Prume et al. 2020). Based on TEM studies also higher mean SV diameters of ∼40 nm and beyond have been published for different species and brain regions (e.g. Schikorski and Stevens 1997; Sätzler et al. 2002; Deak et al. 2003; Hu et al. 2008). TEM tomographic analyses also indicate variability in SV diameter from ∼30 to ∼52 nm in different brain regions and mouse knock-out model systems (Imig et al. 2014; Muth et al. 2024). High variability of SV diameter have been reported to be a general phenomenon in several species and brain regions, even a synapse-to-synapse variability in mean SV size is described in neighboring synapses in the hippocampus, neocortex, and cerebellum of rats, mice, and human cell culture systems (Hu et. al. 2008; Qu et al. 2009). Additionally, within a single SB the size of individual SVs can vary for example depending on their transmitter filling state and content (e.g. Colliver et al. 2000; Pothos et al. 2002).

Size of the three pools of SVs

Besides the size of the PreAZ, the pool of rapidly releasable SVs also critically determines Pr and thus synaptic efficacy, strength and plasticity (Rosenmund and Stevens 1996; Schikorski and Stevens 2001; Rizzoli and Betz 2004; Schikorski 2014; Watanabe et al. 2014; Vaden et al. 2019; reviewed by Rizzoli and Betz 2005, Neher 2015; Chamberland and Toth 2016). It is still rather unclear whether functionally heterogeneous SV pools are structurally identifiable and thus support diverse forms of synaptic transmission und would also play a pivotal role in long- and short-term plasticity. Synaptic transmission can be modulated in various ways depending on the availability of SVs and on their recycling rates. Hence, the size of both the RRP and RP critically determines synaptic efficacy, strength and plasticity. These parameters are controlled at the PreAZ but vary substantially across various CNS synapses (reviewed by Rizzoli and Betz 2005; Neher 2015). The contribution of the RRP size to synaptic dynamics and the mechanisms by which such control is achieved at individual SBs remains largely unknown.

Earlier version: L1 SBs had a total pool size of ∼3500 SVs/AZ, the largest total pool size when compared with other layers of the human TLN by ∼2-fold (L4, ∼1800 SVs) by ∼2.6-fold (L5, ∼1350 SVs) and by ∼3-fold (L6, ∼1150 SVs), respectively. Remarkably, the total pool sizes in the human TLN were significantly larger by more than 6-fold (∼550 SVs/AZ), and ∼4.7-fold (∼750 SVs/AZ;) than those in L4 and L5 (Yakoubi et al. 2019a, b; see also Rollenhagen et al. 2018) in rats. The largest total pool size in L1 may point to a pivotal role in the rapid replenishment after depletion by high-frequency stimulation and the transfer of SVs into the RRP and RP from the total pool in L1 SBs.

Revised version: L1-SBs had a total pool size of ∼3500 SVs/AZ, larger than in other layers of the human TLN: ∼2-fold in L4 (∼1800 SVs; Yakoubi et al. 2019b), ∼2.6-fold in L5 (∼1350 SVs; Yakoubi et al. 2019a) and ∼3-fold in L6 (∼1150 SVs; Schmuhl-Giesen et al. 2022). Remarkably, the total pool sizes in L4 and L5 of the human TLN were significantly larger than those in L4 (∼550 SVs/AZ; Rollenhagen et al. 2015) and L5 (∼750 SVs/AZ; Rollenhagen et al. 2018) of the rat somatosensory cortex. The largest total pool size in L1 may point to a pivotal role in the rapid replenishment after depletion by high-frequency stimulation and the transfer of SVs into the RRP and RP from the total pool in L1 SBs.

Earlier version: The putative RRP at L1 PreAZs was on average ∼4 SVs/AZ for the p10 nm RRP and similar to that in L5 (∼5 SVs/AZ), but significantly smaller by ∼5.3-fold to that of L4 (∼20 SVs/AZ) and by ∼3.5-fold when compared to L6 in the human TLN. Also, huge differences were observed in the p20 nm RRP which at L1 PreAZs was ∼20 SVs, ∼40 SVs in L4, ∼15 SVs in L5 and ∼30 SVs in L6, respectively. It has to be noted that a layer-specific difference in both the p10 nm and p20 nm RRP exists although a huge variability was observed as indicated by the SD, IQR, and variance. Hence, at L1 PreAZs the overall RRP, taken the p10 nm and p20 nm criterion (both together are less than a vesicle diameter) was constituted by ∼25 SVs rapidly available SVs that may partially contribute to the efficacy, strength and reliability of synaptic transmission, but also in the modulation of short-term synaptic plasticity in L1.

Revised version: The putative RRP at L1 PreAZs was on average ∼4 SVs/AZ for the p10 nm RRP and similar to that in L5 (∼5 SVs/AZ), but significantly smaller by ∼5.3-fold to that of L4 (∼20 SVs/AZ) and by ∼3.5-fold when compared to L6 in the human TLN. Also, huge differences were observed in the putative p20 nm RRP which at L1 PreAZs was ∼20 SVs, ∼40 SVs in L4, ∼15 SVs in L5 and ∼30 SVs in L6, respectively. It has to be noted that a layer-specific difference in both the putative p10 nm and p20 nm RRP exists although a huge variability was observed as indicated by the SD, IQR, and variance. Hence, at L1 PreAZs the overall RRP, taken the p10 nm and p20 nm criterion (both together are less than a vesicle diameter) was constituted by ∼25 SVs rapidly available SVs that may partially contribute to the efficacy, strength and reliability of synaptic transmission, but also in the modulation of short-term synaptic plasticity in L1.

Earlier version: The size of the putative RP/PreAZ was ∼470 SVs in human L1 SBs and thus ∼3-fold larger to that in L6 SBs, ∼1.3-fold larger than in L4, and ∼2.4-fold larger than in L5 SBs of the human TLN. In the rodent neocortex the RP/PreAZ comprised ∼130 (L4; see also Rollenhagen et al. 2015) and ∼200 SVs (L5; see also Rollenhagen et al. 2018). The comparably largest size of the RP in L1 when compared with values for the other cortical layers in the human TLN and that found for L4 and L5 in rodents may also point to a rapid availability of SVs from the RP after the replenishment of the RRP during high-frequency stimulation. If the refilling rates were activity dependent, the large size of the RP at L1 PreAZs could explain some forms of short-term synaptic plasticity, e.g. a substantial increase in synaptic strength during frequency facilitation and post-tetanic potentiation at these SBs.

Revised version: The size of the putative RP/PreAZ was ∼470 SVs in human L1 SBs and thus ∼3-fold larger to that in L6 SBs, ∼1.3-fold larger than in L4, and ∼2.4-fold larger than in L5 SBs of the human TLN. In the rodent neocortex the RP/PreAZ comprised ∼130 (L4; Rollenhagen et al. 2015) and ∼200 SVs (L5; Rollenhagen et al. 2018). The comparably largest size of the RP in L1 when compared with values for the other cortical layers in the human TLN and that found for L4 and L5 in rodents may also point to a rapid availability of SVs from the putative RP after the replenishment of the putative RRP during high-frequency stimulation. If the refilling rates were activity dependent, the large size of the RP at L1 PreAZs could explain some forms of short-term synaptic plasticity, e.g. a substantial increase in synaptic strength during frequency facilitation and post-tetanic potentiation at these SBs.

Earlier version: Finally, the resting pool of SVs at L1 PreAZs in the human TLN is also comparably large (∼3000 SVs) and again the largest when compared with other layers in the human TLN (L4 and L5: ∼1250 SVs; L6: ∼900 SVs). The size of the resting pool may guarantee to rapidly replenish the RRP and RP after repetitive high-frequency stimulation via active transfer of SVs with the help of mitochondria associated with the pool of SVs (this study, see also Zhu and Fuster 1996; Verstreken et al. 2005; Smith et al. 2016).

Revised version: Finally, the putative resting pool of SVs at L1 PreAZs in the human TLN is also comparably large (∼3000 SVs) and again the largest when compared with other layers in the human TLN (L4 and L5: ∼1250 SVs; L6: ∼900 SVs). The size of the resting pool may guarantee to rapidly replenish the RRP and RP after repetitive high-frequency stimulation via active transfer of SVs with the help of mitochondria associated with the pool of SVs (this study, see also Zhou and Fuster 1996; Verstreken et al. 2005; Smith et al. 2016). The notable disparities in AZ and SV pool sizes among individual SBs may contribute to rapid alterations in the computational properties of single neurons or networks. Consequently, these variations in AZ and SV pool sizes at L1 SBs may critically influence the behavior of SBs during so-called Up-and-Down states as described for other SBs of the CNS (Zhou and Fuster 1996; Sanchez-Vives and McCormick 2000; Sakata and Harris 2009; Testa-Silva et al. 2014).

Removed in the revised version: The notable disparities in AZ and SV pool sizes among individual SBs may contribute to rapid alterations in the computational properties of single neurons or networks. Consequently, these variations in AZ and SV pool sizes at L1 SBs may critically influence the behavior of SBs during so-called Up-and-Down states as described for other SBs of the CNS (Zhou and Fuster 1996; Sanchez-Vives and McCormick 2000; Sakata and Harris 2009; Testa-Silva et al. 2014).

Importance of presynaptic mitochondria in synaptic transmission

Mitochondria in the cortical layers of the human TLN were organized in clusters associated with the pool of SVs (this study, see also Yakoubi et al. 2019a, b; Schmuhl-Giesen et al. 2022). In L1, mitochondria only contribute by about 7.2% (L1a) and 6.7% (L1b) to the total SB volume with similar values between L1 and L6, but their percentage was ∼2-fold lower than values in L4 and L5 suggesting a layer-specific distribution of mitochondria in the human TLN.

Earlier version: Mitochondria play a role in the recruitment and mobilization of SVs from the RP and resting pool and in the priming and docking process of SVs (Verstreken et al. 2005; Perkins et al. 2010; Smith et al. 2016; reviewed by Dallerac et al. 2018). In the CNS they act as the main source of internal calcium (Pozzan and Rizzuto 2000; Rizzuto et al. 2000), thus they regulate and control the internal calcium concentrations in CNS terminals required for the signal cascades where for example, synaptic proteins driven by Ca2+ like synaptotagmin, synaptophysin, synaptobrevin and the SNARE-complex are involved (Südhof 2002).

Revised version: Mitochondria play a pivotal role in the recruitment and mobilization of SVs from the RP and resting pool and in the priming and docking process of SVs (Verstreken et al. 2005; Perkins et al. 2010; Smith et al. 2016; reviewed by Dallerac et al. 2018). In the CNS they act as the main source of internal calcium (Pozzan and Rizzuto 2000; Rizzuto et al. 2000), thus they regulate and control the internal calcium concentrations in CNS terminals required for the signal cascades where for example, synaptic proteins driven by Ca2+ like synaptotagmin, synaptophysin, synaptobrevin and the SNARE-complex are involved (Südhof 2002).

Astrocytic coverage of L1 SBS in the human TLN

Earlier version: It is widely recognized that astrocytes play a crucial role in the formation of the ‘tripartite’ synapse, which is a common characteristic of cortical synapses. Astrocytes serve as both a physical barrier to glutamate diffusion and as mediate neurotransmitter uptake via transporters, thereby regulating the spatial and temporal concentration of neurotransmitters in the synaptic cleft (Oliet et al. 2004; Min and Nevian 2012; Pannasch et al. 2014; reviewed by Dallerac et al. 2018). In addition, they modulate synaptic transmission by activating pre- and postsynaptic receptors (Haydon and Carmignoto 2006; Le Meur et al. 2012). Moreover, it was found that the control of timing-dependent long-term depression (t-LTD) at neocortical synapses is critically influenced by astrocytes by increasing Ca2+ signaling during the induction of t-LTD (Min and Nevian 2012).

Revised version: It is widely recognized that astrocytes play a crucial role in the formation of the ‘tripartite’ synapse, which is a common characteristic of cortical synapses. Astrocytes serve both as a physical barrier to glutamate diffusion and to mediate neurotransmitter uptake via transporters, thereby regulating the spatial and temporal concentration of neurotransmitters in the synaptic cleft (Oliet et al. 2004; Min and Nevian 2012; Pannasch et al. 2014; reviewed by Dallerac et al. 2018). In addition, they modulate synaptic transmission by activating pre- and postsynaptic receptors (Haydon and Carmignoto 2006; Le Meur et al. 2012). Moreover, it was found that the control of timing-dependent long-term depression (t-LTD) at neocortical synapses is critically influenced by astrocytes by increasing Ca2+ signaling during the induction of t-LTD (Min and Nevian 2012).

Remarkably, significant layer-specific differences exist in the astrocytic coverage of synaptic complexes in the human TLN. Whereas in L4 and L5 ∼60-80% of the total volumetric fraction was occupied by astrocytic processes and most of AZs in L4 were tightly ensheathed by fine astrocytic processes, in L1 and L6 the volumetric fraction of astrocytic processes was only ∼20% with a loose and incomplete coverage of AZs in L1 (Figure 6D). As a consequence, the lack or only incomplete astrocytic coverage of synaptic complexes in L1 can only partially act as a physical barrier to neurotransmitter diffusion. In addition, beside a vertical also its horizontal diffusion of glutamate at the synaptic cleft is possible. Hence, two possible scenarios for the role of astrocytes are present at L1 synaptic complexes. First, as shown for only a small part (10%) of synaptic complexes in L1 the ‘tripartite’ synapse is realized at both sides of the synaptic cleft (Figure 6D). This would allow the selective uptake of horizontally diffusing glutamate via glutamate transporters located in the fine astrocytic processes. This prevents ‘glutamate spillover’, and thus allow a vertically directed neurotransmitter diffusion and docking to postsynaptic receptors at the PSD.

The remaining larger part of synaptic complexes in L1 either lack a coverage by fine astrocytic processes at AZs or are only partially ensheathed (Figure 6D). Here ‘glutamate spillover’ of horizontally diffusing neurotransmitter quanta is most likely and thus synaptic crosstalk between neighboring synaptic complexes at AZs in neighboring spines on terminal tuft dendrites because spine density in L1 is high. This may cause a switch from asynchronous to synchronous release from neighboring synaptic complexes upon repetitive low- and high-frequency stimulation (von Gersdorff and Borst 2002).

In summary, two distinct structural scenarios were identified regarding the role of astrocytes in L1, However, it is still rather unknown which of the two scenarios contribute more efficiently to synaptic transmission and plasticity in L1.

The role of L1 in the information processing of the cortical column

Earlier version: The computational properties of the neocortex depend upon its ability to integrate the information provided by the sensory organs (bottom-up information) with internally generated signals such as expectations or attentional signals (top-down information). This integration occurs at apical tuft dendrites of L2/3 and L5 pyramidal neurons in L1. Importantly, L1 is the predominant input layer for top-down information, relayed by a rich, dense network of cortico-cortical, commissural and associational long-range axonal projections and L1 excitatory and inhibitory neurons providing signals to the terminal tuft dendrites of pyramidal neurons (reviewed by Schumann et al. 2021). Thus, L1 is a central locus of neocortical associations controlled by distinct types of excitation and inhibition in humans (reviewed by Hartung and Letzkus 2021).

Revised version: The computational properties of the neocortex depend upon its ability to integrate the information provided by the sensory organs (bottom-up information) with internally generated signals such as expectations or attentional signals (top-down information). This integration occurs at apical tuft dendrites of L2/3 and L5 pyramidal neurons in L1. Importantly, L1 is the predominant input layer for top-down information, relayed by a rich, dense network of cortico-cortical, commissural and associational long-range axonal projections. In experimental animals, L1 consists exclusively of a heterogeneous population of GABAergic interneurons providing feedforward and feedback inhibition (see for example Jiang et al. 2013; Boldog et al. 2018; Obermayer et al. 2018; Kwon et al. 2019; reviewed by Huang et al. 2024) providing signals to the terminal tuft dendrites of pyramidal neurons (reviewed by Schumann et al. 2021). In humans, there is a persistent subpopulation of CR cells that provides excitation due to their long-range transcolumnar axonal projection to the terminal tuft dendrites (Anstötz et al. 2014), but may also control the population of GABAergic interneurons, which is as to date unknown.

Removed in the revised version: From L4, as the first station of intracortical information processing thalamocortical signals are transferred via intracolumnar projecting axonal collaterals via L2, L3 to L1 (Feldmeyer et al. 1999; Lübke et al. 2000) representing a convergent input layer for thalamocortical (Rodriguez-Moreno et al. 2018) and cortico-cortical synaptic inputs. Hence, L4 excitatory neurons act as ‘feed-back amplifiers’, but also guarantee a highly reliable signal transduction to L3, L2 and L1 (Feldmeyer et al. 1999; Seeman et al. 2018).

Removed in the revised version: SBs in L1 contact both the persistent subpopulation of CR-cells, various GABAergic interneurons and terminal tuft dendrites of pyramidal neurons (Figure 7) proving feed-forward excitation to their putative partners in L1. CR-cells exhibit long-range excitation that preferentially and heavily contact terminal tuft dendrites of pyramidal neurons within and across cortical columns via long horizontal axonal collaterals (Anstötz et al. 2014). As shown in rodent neocortex, GABAergic interneurons are synaptically coupled with CR-cells providing feed-forward inhibition (Anstötz et al. 2014) thus controlling the output of CR-cells.

Removed in the revised version: Schematic drawing of the possible wiring of L1 in the human TLN.

Removed in the revised version: The possible neuronal and synaptic organization of neocortical L1 in humans modified and based on findings from rat neocortex (reviewed by Schuman et al. 2021). The dendritic domain of excitatory neurons is given in blue, the axonal projection and connectivity in magenta and that of different GABAergic interneurons in red and green, respectively. Furthermore, a subpopulation of persistent CR-cells is illustrated in the diagram. Abbreviations: α7: nicotinic receptor-expressing interneuron; CanC: canopy cell; ChC: chandelier cell; L: layer; NGFC: neurogliaform cell; PC: pyramidal cell; SST: somatostatin containing interneuron; SStC: spiny stellate cells; VIP: vasoactive intestinal peptide-expressing interneuron; VIP/Bip: vasoactive intestinal peptide–expressing bipolar interneuron; VIP/CCK: cholecystokinin- and vasoactive intestinal peptide-expressing interneuron.

Removed in the revised version: In general, GABAergic inhibition plays a critical role in shaping neuronal activity in the neocortex. However, most inhibitory synapses in the neocortex are formed onto pyramidal cell dendrites suggesting they may focally regulate cellular activity (Kwon et al. 2018). Somatostatin-expressing interneurons in L1 exert compartmentalized control over postsynaptic Ca2+ signals mediated by a subset of GABAergic synapses that directly target spine heads. GABAergic inhibition thus participates in localized control of dendritic electrical and biochemical signaling (Boldog et al. 2018). Furthermore, another study has shown that a surprisingly large fraction (33.4 ± 9.3%) of terminals formed symmetrical synaptic junctions onto a subset of cortical spines that were mostly co-innervated by an asymmetrical terminal. The double-innervated spines selectively received thalamocortical afferents expressing VGLUT2 but were not innervated by intracortical synaptic contacts expressing VGLUT1 (Kubota et al. 2007). These results demonstrate that inhibitory inputs to pyramidal cell spines may preferentially reduce thalamocortical rather than intracortical synaptic transmission and are therefore positioned to selectively gate extracortical information.

Removed in the revised version: Another subpopulation of GABAergic interneurons in L1 exhibit a highly specialized form of recurrent inhibition that link cortical L1 to L3 interneurons and L5 pyramidal neurons (Jiang et al. 2013). L1 single-bouquet cells preferentially form inhibitory connections on L2/3 interneurons that inhibit the entire dendritic-somato-axonal axis of L5 pyramidal neurons located within the same column. In contrast, L1 neurogliaform cells frequently form inhibitory and electric connections with L2/3 interneurons. These L1-L3 interneurons inhibit the distal apical dendrite of > 60% of L5 pyramidal neurons across multiple columns. Functionally, the single bouquet cell to L2/3 interneurons to L5 pyramidal neuronal circuits disinhibit, and neurogliaform cells to L2/3 interneuron to L5 pyramidal neuronal circuits inhibit the initiation of dendritic complex spikes in L5 pyramidal neurons. As dendritic complex spikes can serve coincidence detection, these cortical interneuronal circuits may be essential for salience selection (Jiang et al. 2013).

Removed in the revised version: Finally, L1 receive GABAergic input from Martinotti cells located in L2/3 and L5, axo-axonic Chandelier cells in L2/3 and VIP-expressing interneurons in L2 all of which have a very prominent axonal plexus in L1 (Figure 7, see also Silberberg and Markram 2007; Obermayer et al. 2018; reviewed by Hartung and Letzkus 2021; Schuman et al. 2021; reviewed by Huang et al. 2024).

Earlier version: In summary, these GABAergic interneuron populations provide inhibition and disinhibition to pyramidal terminal tuft dendrites by either shunting or promoting the generation of Ca2+- and Na+-spikes by temporal and spatial coincidence detection. Thus, the circuitry of L1, including the persistent subpopulation of excitatory CR-cells; is an important and fundamental ‘integrator’ of the neocortex sampling bottom-up information within or even across cortical columns, but also act as a ‘filter’ or ‘discriminator’ but also as an ‘amplifier’ in the information processing of the neocortex.

Revised version: Thus, L1 in the human neocortex is a central locus of neocortical associations controlled by excitation and distinct types of inhibition (reviewed by Hartung and Letzkus 2021; Schumann et al. 2021). In summary, the population of GABAergic interneurons in L1 provides inhibition and disinhibition to pyramidal terminal tuft dendrites by either shunting or promoting the generation of Ca2+- and Na+-spikes by temporal and spatial coincidence detection. Thus, the circuitry of L1, including the persistent subpopulation of excitatory CR-cells; is an important and fundamental ‘integrator’ of the neocortex sampling bottom-up information within or even across cortical columns, but also act as a ‘filter’ or ‘discriminator’ but also as an ‘amplifier’ in the information processing of the neocortex.

Removed in the revised version: Hence, L1 of the neocortex can be regarded as a complex and intricate layer of the brain serving different functions, for example to gate inputs, convey expectations and context as well as mediate states of consciousness, attention, cross-modal interactions, sensory perception, and learning (Bastos et al. 2012; Heeger 2017; Zagha 2020; reviewed by Gilbert and Li 2013). There is evidence that L1 processing, particularly Ca2+ signals in the distal apical dendrite, is important for sensory perception (Xu et al. 2012).

Earlier version: Recent studies have revealed that L1 plays a crucial role in top-down attentional processes, aiding in directing attention, based on our goals and intentions (Hartung and Letzkus 2021; Schuman et al. 2021). Additionally, L1 is implicated in the formation and consolidation of long-term memories with disruptions in L1 function associated with memory deficits. Furthermore, L1 contributes to brain plasticity, which involves the brain’s ability to change and adapt in response to experience, including the process by which synapses alter in strength and number (reviewed by Hartung and Letzkus 2021; Schuman et al. 2021).

Revised version: Hence, L1 of the neocortex can be regarded as a complex and intricate layer of the brain serving different functions, for example to gate inputs, convey expectations and context as well as mediate states of consciousness, attention, cross-modal interactions, sensory perception, and learning (Bastos et al. 2012; Heeger 2017; Zagha 2020; reviewed by Gilbert and Li 2013). There is evidence that L1 processing, particularly Ca2+ signals in the distal apical dendrite, is important for sensory perception (Xu et al. 2012). Recent studies have revealed that L1 plays a crucial role in top-down attentional processes, aiding in directing attention, based on our goals and intentions (Hartung and Letzkus 2021; Schuman et al. 2021). Additionally, L1 is implicated in the formation and consolidation of long-term memories with disruptions in L1 function associated with memory deficits. Furthermore, L1 contributes to brain plasticity, which involves the brain’s ability to change and adapt in response to experience, including the process by which synapses alter in strength and number (reviewed by Hartung and Letzkus 2021; Schuman et al. 2021).

Overall, L1 of the neocortex is a complex and fascinating layer of the brain that is involved in a wide range of cognitive processes and brain plasticity. Since this is relayed by a network of SBs in L1, their composition and connections play a pivotal role, in columnar information processing.

Material and methods

Human brain tissue sampling during epilepsy surgery have been provided by Dr. med. Dorothea Miller, PD Dr. med. Marec von Lehe, Department of Neurosurgery, Knappschaftschafts/Universitäts-Krankenhaus Bochum and were approved by the Ethical Committees of the Rheinische Friedrich-Wilhelms-University/University Hospital Bonn (Ethical votum of the Medical Faculty to Prof. Dr. med. Johannes Schramm and Prof. Dr. rer. nat. Joachim Lübke, Nr. 146/11), the University of Bochum (Ethical votum of the Medical Faculty to PD Dr. med. Marec von Lehe and Prof. Dr. rer. nat. Joachim Lübke, Reg. No. 5190-14-15; and renewed Ethical votum of the Medical Faculty to Dr. med. Dorothea Miller and Prof. Dr. rer. nat. Joachim Lübke, Reg. No. 17–6199-BR). The consent of the patients was obtained by written and signed statements, and all further experimental procedures were approved by the same Ethical Committees cited above, and the EU directive (2015/565/EC and 2015/566/EC) concerning working with human tissue used for experimental and scientific purposes. All subsequent experimental procedures were approved by the Research Committee of the Research Centre Jülich GmbH.

Fixation and tissue processing for TEM

Earlier version: Tissue samples from the human TLN were after their removal prepared and embedded for conventional TEM and EM tomography analysis. All neocortical access tissues were obtained from patients (1 male and 3 female, 24-65 years in age, see Supplemental Table 2) suffering from drug-resistant temporal lobe epilepsy. The pre-surgical work-up comprised at least high-resolution magnetic resonance imaging together with long-term video-electro-encephalography (EEG)-monitoring. In all cases, the circumscribed epileptic focus was in the hippocampus proper, but not in the neocortical access regions of the TL.

Revised version: Tissue samples from the human TLN were after their removal prepared and embedded for conventional TEM and EM tomography analysis. All neocortical access tissues were obtained from patients suffering from drug-resistant temporal lobe epilepsy (1 male and 3 female, 24-65 years in age for L1, and additionally 2 male, and 3 female, 25-63 in age for L4; see also Supplemental Table 2).

Added in the revised version: The pre-surgical work-up comprised high-resolution magnetic resonance imaging together with long-term video-electro-encephalography (EEG)-monitoring and electrophysiology using multiple electrode recordings. In all cases, the circumscribed epileptic focus was in the hippocampus proper, but not in the neocortical access regions of the TL.

Added in the revised version: During epilepsy surgery, blocks of both non-affected, non-epileptic and epileptic neocortical access tissue samples (see Supplemental Table 2) were resected and part of the tissue samples were further histologically examined by neuropathologists. The non-affected non-epileptic neocortical access tissue samples were always taken far from the epileptic focus and may thus be regarded as non-epileptic as also demonstrated by functional studies using the same experimental approach (Mohan et al. 2015; Molnár et al. 2016; Seeman et al. 2018; Obermayer et al. 2018; reviewed by Mansvelder et al. 2019). These studies have clearly demonstrated that neocortical access tissue samples taken from epilepsy surgery do not differ in electrophysiological properties and synaptic physiology when compared with acute slice preparations in experimental animals.

Earlier version: During epilepsy surgery, blocks of both non-affected, non-epileptic and epileptic neocortical access tissue samples (see Supplemental Table 2) were resected to control the seizures for histological inspection by neuropathologists. The non-affected non-epileptic neocortical access tissue samples was always taken far from the epileptic focus and may thus be regarded as non-epileptic as also demonstrated by other structural and functional studies using the same experimental approach (Alonso-Nanclares 2008; Navarette et al. 2013; Mohan et al. 2015; Molnár et al. 2016, Seeman et al. 2018; reviewed by Mansfelder et al. 2019). So-called ‘post mortem’ tissue was not used in this study since the ultrastructural quality (preservation) of such material is not suitable enough for fine-scale high-resolution EM, due to severe distortions of relevant structural features, e.g. fragmentation and lysis of membranes of SBs, PreAZs, PSDs and SVs that are required for the generation of quantitative 3D-models of SBs and their target structures (Lübke, personal observation). However, under certain conditions also ‘post mortem’ tissue samples can be used (Dominguez-Alvaro et al. 2019, 2921; Cano-Astorga et al. 2021, 2023).

Revised version: So-called ‘post mortem’ tissue was not used in this study since the ultrastructural quality (preservation) of such material is not suitable enough for fine-scale high-resolution EM, due to severe distortions of relevant structural features, e.g. fragmentation and lysis of membranes of SBs, PreAZs, PSDs and SVs that are required for the generation of quantitative 3D-models of SBs and their target structures (Lübke, personal observation). However, under certain conditions also ‘post mortem’ tissue samples can be used (Dominguez-Alvaro et al. 2019, 2921; Cano-Astorga et al. 2021, 2023).

Removed in the revised version: Inter-individual differences were found in the structural and synaptic parameters analyzed as shown by the box plots (Supplemental Figures 1, 2; Source Data 02).

Earlier version: Other recent studies using the same experimental approach neglected and thus discarded the effect of the pharmacological treatments and disease condition (Alonso-Nanclares et al. 2008; Testa Silva et al. 2014; Mohan et al. 2015; Molnár et al. 2016; Seeman et al. 2018; Dominguez-Alvaro et al. 2019; Yakoubi et al. 2019a, b; Schmuhl-Giesen et al. 2022).

Revised version: Inter-individual differences were found in the structural and synaptic parameters analyzed as shown by the box plots (Supplemental Figures 1, 2; Source Data 02). Other recent studies using the same experimental approach neglected and thus discarded the effect of the pharmacological treatments and disease condition (Alonso-Nanclares et al. 2008; Testa Silva et al. 2014; Mohan et al. 2015; Molnár et al. 2015; Seeman et al. 2018; Dominguez-Alvaro et al. 2019; Yakoubi et al. 2019a, b; Schmuhl-Giesen et al. 2022).

For the study of L1 in the human TLN, blocks of neocortical access tissue were sampled from the temporo-lateral or temporo-basal regions of the inferior temporal gyrus and the gyrus medialis. Immediately after their removal during epilepsy surgery, biopsy samples of the TLN were immersion-fixed in ice-cold 4% paraformaldehyde and 2.5% glutaraldehyde diluted in 0.1 M phosphate buffer (PB, pH 7.4) for 24-72 hrs at 4°C. The fixative was replaced by fresh fixative after 2 hrs and changed twice during the subsequent fixation period. Prior to vibratome sectioning, brain tissue samples were thoroughly rinsed in ice-cold PB and afterwards embedded in 5% Agar-Agar (Sigma, Munich, Germany) diluted in PB.

Neocortical tissue blocks were cut in the coronal plane through the TLN with a Vibratome VT 1000S (Leica Microsystems GmbH, Wetzlar, Germany) into 150-200 µm thick sections, collected in ice-cold PB, and washed again several times in PB. Afterwards they were transferred to 0.5-1% PB-sucrose buffered Osmium tetroxide (OsO4, 300 mOsm, pH 7.4; Sigma, Munich, Germany) for 60-90 min. After visual inspection to check for the quality of post-fixation, sections were thoroughly washed several times in PB and left overnight at 4°C in PB. The next day they were dehydrated in an ascending series of ethanol starting at 20%, 30%, 50%, 60%, 70%, 80%, 90%, 95% to absolute ethanol (15 min for each step and absolute ethanol, 30 min twice), followed by a brief incubation in propylene oxide (2 min twice; Fluka, Neu-Ulm, Germany). Sections were then transferred into a mixture of propylene oxide and DurcupanTM resin (2:1, 1:1 for 1h each; Fluka, Neu-Ulm, Germany) and stored overnight in pure resin. The next day, sections were flat embedded on coated glass slides in fresh DurcupanTM, coverslipped with AclaTM foils and polymerized at 60°C for 2 days.

Semi- and ultrathin sectioning

After light microscopic (LM) inspection, a tissue block containing the region of interest (ROI) was glued on a pre-polymerized block and trimmed down. Semithin sections were cut with a Leica UltracutS ultramicrotome (Leica Microsystems, Vienna, Austria), with a Histo-Diamond knife (Fa. Diatome, Nidau, Switzerland). Afterwards they were briefly stained with methylene-blue (Sigma-Aldrich Chemie GmbH, Taufkirchen, Germany) to identify the cortical layers, particularly L1 and underlying L2, examined and documented using a motorized Olympus BX61 microscope equipped with the Olympus CellSense analysis hard- and software (Olympus GmbH, Hamburg, Germany). All images were stored in a database until further use.

After further trimming of the block to its final size, serial ultrathin sections (50 ± 5 nm thickness; silver to silver gray interference contrast) were cut with a Leica UltracutS ultramicrotome through the determined ROI of L1a and L1b, respectively. Ultrathin sections in serial sequence were collected on pioloform-coated slot copper grids (Plano, Wetzlar, Germany). Prior to EM examination, sections were stained with 5% aqueous uranyl acetate or Uranyless (Science Services, Munich, Germany) for 15–20 min and lead citrate for 3–5 min according to Reynolds (1963) to enhance the contrast of biological membranes at the EM level. In each series of ultrathin sections, two or three ROIs throughout L1a or L1b were chosen, and photographed at a final EM magnification of 8000x with a Zeiss Libra 120 (Fa. Zeiss, Oberkochen, Germany) equipped with a Proscan 2K digital camera (Fa. Tröndle, Moorenweis, Germany) using the ImageSP software and its panorama function (Fa. Tröndle, Moorenweis, Germany). In addition, interesting details of synaptic structures in L1a and L1b were taken at various EM magnifications.

All images were stored in a database until further use. Selected EM images for publication were further edited using Adobe PhotoshopTM and Adobe IllustratorTM software.

3D-volume reconstructions and quantitative analysis of SBs

EM panorama images composing each series were imported, stacked, and aligned in the reconstruction software OpenCAR (Contour Alignment Reconstruction; for details see Sätzler et al. 2002). The main goal of this study was to quantify several structural and synaptic parameters representing structural correlates of synaptic transmission and plasticity. Excitatory SBs were characterized by large round SVs and prominent PreAZs and PSDs in contrast to putative GABAergic terminals that have smaller, more oval-shaped SVs, and a thin or no PSDs.

Earlier version: The following structural parameters were analyzed: (1) surface area and volume of SBs; (2) volume of mitochondria; (3) surface area of PreAZs (Dufour et al. 2016) and PSDs; two apposed membrane specializations separated by the synaptic cleft; (4) number and diameter of clear SVs and dense-core vesicles (DCVs); and (5) distance of individual SVs from the PreAZ for the structural definition of the RRP, RP, and resting pool. 3D-volume reconstructions were generated by drawing contour lines on the structures of interest in each panorama image within a series of images until a given structure was completed. Presynaptic SBs, their mitochondria as well as their postsynaptic target structures were outlined on the outer edge of their membranes, using closed contour lines throughout the entire stack within a series. A SB was considered completely captured, when it was possible to follow the axon in both directions through the entire series (en passant SBs) or the enlargement of the axon leading to an endterminal SB. The beginning of a synaptic terminal was defined by the typical widening of the axon and the abrupt occurrence of a pool of SVs.

Revised version: The following structural parameters were analyzed in randomly selected excitatory SBs: (1) surface area and volume of SBs; (2) volume of mitochondria; (3) surface area of PreAZs (Dufour et al. 2016) and PSDs; two apposed membrane specializations separated by the synaptic cleft; (4) number and diameter of clear SVs and dense-core vesicles (DCVs); and (5) distance of individual SVs from the PreAZ for the structural definition of the RRP, RP, and resting pool. 3D-volume reconstructions were generated by drawing contour lines on the structures of interest in each panorama image within a series of images until a given structure was completed. Presynaptic SBs, their mitochondria as well as their postsynaptic target structures were outlined on the outer edge of their membranes, using closed contour lines throughout the entire stack within a series. A SB was considered completely captured, when it was possible to follow the axon in both directions through the entire series (en passant SBs) or the enlargement of the axon leading to an endterminal SB. The beginning of a synaptic terminal was defined by the typical widening of the axon and the abrupt occurrence of a pool of SVs.

Earlier version: The PreAZs and PSDs were regarded as complete when their perimeters were entirely reconstructed in a series of ultrathin sections. The surface areas of the PreAZ and PSD were computed separately by first generating a 3D surface model of the SB. The PreAZ was then measured by extracting this area from the reconstructed presynaptic bouton membrane that was covered by this membrane specialization. Hence, the length (L) of the PreAZ (L PreAZ) and the surface area (SA) of the PreAZ (SA PreAZ) is already known.

Revised version: The PreAZs and PSDs were regarded as complete when their perimeters were entirely reconstructed in a series of ultrathin sections. The surface areas of the PreAZ and PSD were computed separately by first generating a 3D surface model of the SB. The PreAZ was then measured by extracting this area from the reconstructed presynaptic bouton membrane that was covered by this membrane specialization. Hence, the length (L) of the PreAZ (L PreAZ) and the surface area (SA) of the PreAZ (SA PreAZ) is already known. The size of the PSD opposing the PreAZ was estimated under the following assumptions: (1) both membrane specializations, PreAZ and PSD run parallel to each other at the pre- and postsynaptic apposition zone; (2) for both membrane specializations a contour line was drawn determining their actual length (L PreAZ and L PSD). Hence, the surface area of the PSD (SA PSD) is estimated by the following equation: which is the perimeter ratio between the outlines of the PSD to that of the synaptic contact.

Removed in the revised version: The size of the PSD opposing the PreAZ was estimated under the following assumptions: (1) both membrane specializations, PreAZ and PSD run parallel to each other at the pre- and postsynaptic apposition zone; (2) for both membrane specializations a contour line was drawn determining their actual length (L PreAZ and L PSD). Hence, the surface area of the PSD (SA PSD) is estimated by the following equation:

Removed in the revised version: which is the perimeter ratio between the outlines of the PSD to that of the synaptic contact.

The synaptic cleft width was measured because of its importance for the transient temporal and spatial increase of the glutamate concentration, reversible binding of glutamate to appropriate glutamate receptors and eventual up-take and diffusion of glutamate out of the synaptic cleft. To a large extent, these processes are governed by the geometry of this structure and the shape and size of the PreAZs and PSDs.

All SVs were marked throughout each SB, their diameters were individually measured and their distances to the PreAZ were automatically detected using an algorithm implemented in OpenCAR (minimal distance between each SV membrane to the contour line of the PreAZ) throughout the SB in every single image of the series. Large DCVs were only counted in the image where they appeared largest. To avoid double counts, only clear ring-like structures were counted as SVs. However, SVs might be missed in densely packed regions, because ring-like traces may partly overlap. This effect may counteract any double counts. Based on the small extent and the partially counteracting nature of this effect, the numbers of small clear vesicles reported in this study remained uncorrected (Yakoubi et al. 2019a, b, Schmuhl-Giesen et al. 2022).

In this work, aldehyde fixation was used that is thought to induce tissue shrinkage thereby biasing structural quantification (Eyre et al. 2007; Korogod et al. 2015). A direct comparison of structural parameters obtained from either aldehyde or cryo-fixed and substituted tissue samples (Korogod et al. 2015), showed differences in cortical thickness (∼16% larger in cryo-fixed material), in the volume of extracellular space (∼6-fold larger in cryo-fixed material), a slight increase in glial volume and overall density of synaptic contacts (∼14% in cryo-fixed material), but no significant differences in neuronal structures such as axons, dendrites and SV diameter.

In the structural and synaptic parameters as estimated here, no significant difference was found for SB size and other synaptic subelements such as mitochondria, active zones (AZs) and SVs when compared with other studies (Zhao et al. 2012a, b). Therefore, no corrections for shrinkage were applied and we are thus convinced that the synaptic parameters reported here are accurate and can be directly used for detailed computational models. In addition, large-scale preservation for ultrastructural analysis will therefore continue to rely on chemical fixation approaches, due to the limited preservation of the ultrastructure in cryo-fixed material as stated in Korogod and co-workers (2015).

Golgi-Cox impregnation of biopsy material in L4 of the human TLN

To get an impression about the neuronal organization of the human TLN, four tissue blocks were processed with the Golgi-Cox impregnation technique using the commercially available Hito Golgi-Cox OptimStain kit (Hitobiotec Corp, Kingsport, TE, USA). After removal of the biopsy samples, tissues were briefly rinsed twice in double distilled water (dd H2O), and then transferred into impregnation solution 1 overnight at room temperature. The next day, tissue samples were incubated in fresh impregnation solution 1 and stored for 14 days in the dark at room temperature. The samples blocks were then transferred in solution 3 and kept in the dark at room temperature for one day. Thereafter, they were placed into fresh solution 3 in the dark at room temperature for 6 additional days. Then, solution 3 was exchanged and samples were stored at 4°C in the dark overnight. Tissue blocks were embedded in 5% agarose (Carl Roth, Karlsruhe, Germany) diluted in ddH2O, and sectioned with a vibratome in the coronal plane at 100–250 µm thickness and then transferred to ddH2O.

After careful removal of the agarose, free-floating sections were incubated into solution 3 for 2–3 min in the dark at room temperature, and right after placed into ddH2O, washed several times and stored overnight. Afterwards, they were rinsed twice in ddH2O for 4 min each, and dehydrated in 50%, 70% and 95% ethanol for 5 min each, then transferred into absolute ethanol (3 × 5 min), defatted in xylene, embedded in EukittTM (Sigma-Aldrich Chemie GmbH, Taufkirchen, Germany), finally coverslipped and air-dried. Afterwards, sections were examined and imaged with an Olympus BX 61 LM equipped with the CellSense software package (Olympus, Hamburg, Germany) at various magnifications and images were stored in a database until further use.

Stereological estimation of the density of L1 synaptic contacts in the human TLN

Earlier version: The density of synaptic complexes, composed either between an SB with a dendrite or spine, in a given volume is a valuable parameter to assess the structural and functional changes in the brain, which are linked to the age, pathological or experimental conditions (Rakic et al. 1994; DeFelipe et al. 1999). The density of synaptic contacts was unbiasedly estimated in L1, separated for L1a and L1b, from four patients, respectively (Supplemental Table 2; Source Data 01) using the physical dissector technique (Mayhew 1996; Fiala and Harris 2001) by counting the synaptic complexes in a virtual volume generated by two adjacent ultrathin sections that is the dissector: the reference section and the look-up section. Here, counting was performed using FIJI (Schindelin et al. 2012) on a stack of 20 aligned serial electron micrographs for each patient taken from the series of ultrathin sections used for the 3D-volume reconstructions of SBs in L1. An unbiased counting frame was first set and synaptic contacts to be considered (counted) are the one present in the reference section only and meeting the following criteria: presence of a PreAZ and a prominent or thin PSD separated by a synaptic cleft and SVs in the presynaptic terminal. Care was taken to distinguish between excitatory and inhibitory synaptic contacts, as well as the postsynaptic target structures (dendritic shafts or spines). Finally, the density of synaptic contacts (Nv) per 1 mm3 was calculated using the formula below:

Revised version: The density of synaptic complexes, composed either between an SB with a dendrite or spine, in a given volume is a valuable parameter to assess the structural and functional changes in the brain, which are linked to the age, pathological or experimental conditions (Rakic et al. 1994; DeFelipe et al. 1999). The density of synaptic contacts was unbiasedly estimated in L1, separated for L1a and L1b, from four patients, respectively (Supplemental Table 2; Source Data 01) using the physical dissector technique (Mayhew 1996; Fiala and Harris 2001) by counting the synaptic complexes in a virtual volume generated by two adjacent ultrathin sections that is the dissector: the reference section and the look-up section. Here, counting was performed using FIJI (Schindelin et al. 2012) on a stack of 20 aligned serial electron micrographs for each patient taken from the series of ultrathin sections used for the 3D-volume reconstructions of SBs in L1. An unbiased counting frame was first set and synaptic contacts to be considered (counted) are the one present in the reference section only and meeting the following criteria: presence of a PreAZ and a prominent or thin PSD separated by a synaptic cleft and SVs in the presynaptic terminal. Care was taken to distinguish between excitatory and inhibitory synaptic contacts, as well as the postsynaptic target structures (dendritic shafts or spines). Finally, the density of synaptic contacts (Nv) per 1 mm3 was calculated using the formula below: where Qd is the number of synaptic contacts per dissector and Vd is the volume of the dissector given by: Number of dissectors x frame area x section thickness.

Removed in the revised version: where Qd is the number of synaptic contacts per dissector and Vd is the volume of the dissector given by: Number of dissectors x frame area x section thickness.

TEM tomography of L1 SBs in the human TLN

Removed in the revised version: EM tomography was performed on 200-300 nm thick sections cut from blocks prepared for serial ultrathin sectioning as described above (Table 2; Source Data 03). Sections were mounted on pioloform-coated line copper grids and were counterstained with uranyl acetate and lead citrate following a slightly modified staining protocol as described by Reynolds (1963).

Earlier version: Subsequently, sections were examined with a JEOL JEM 1400Plus, operating at 120 kV and equipped with a 4096×4096 pixels CMOS camera (TemCam-F416, TVIPS, Gauting, Germany). Tilt-series were acquired automatically over an angular range of −60° to +60° at 1°-degree increments using Serial EM (Mastronarde 2005). Stack alignment and reconstruction by filtered backprojection were carried out using the software package iMOD (Kremer et al. 1963). Final reconstructions were ultimately filtered using a median filter with a window size of 3 pixels. Tilt-series were stored as .tif files and were further processed using the freely available software Imod 4.9.12 and ImageJ (ImageJ, RRID:SCR-0033070). In each tilt-series so-called ‘docked’ SVs identified by their fusion with the PreAZ or as omega-shaped body that already released a quantum of neurotransmitter were counted separately for L1a and L1b and for their target structures, dendritic shafts or spines.

Revised version: EM tomography was performed on 200-300 nm thick sections cut from blocks prepared for serial ultrathin sectioning as described above (Table 2; Source Data 03). Sections were mounted on pioloform-coated line copper grids and were counterstained with uranyl acetate and lead citrate following a slightly modified staining protocol as described by Reynolds (1963). Subsequently, sections were examined with a JEOL JEM 1400Plus, operating at 120 kV and equipped with a 4096×4096 pixels CMOS camera (TemCam-F416, TVIPS, Gauting, Germany). Tilt-series were acquired automatically over an angular range of −60° to +60° at 1°-degree increments using Serial EM (Mastronarde 2005). Stack alignment and reconstruction by filtered backprojection were carried out using the software package iMOD (Kremer et al. 1996). Final reconstructions were ultimately filtered using a median filter with a window size of 3 pixels. Tilt-series were stored as .tif files and were further processed using the freely available software Imod 4.9.12 and ImageJ (ImageJ, RRID:SCR-0033070). In each tilt-series so-called ‘docked’ SVs identified by their fusion with the PreAZ or as omega-shaped body that already released a quantum of neurotransmitter were counted separately for L1a and L1b and for their target structures, dendritic shafts or spines in randomly selected SBs.

Added in the revised version: To determine possible differences to the values obtained by the 3D-reconstructions, in the same tilt-series binned at 24 nm (section thickness) the diameter of SVs was measured randomly by using OpenCAR as described above. To avoid double counts and underestimations of the diameter, the ring-shaped structure of an SV was compared on each tilted-slice and only marked if it appears largest. In 18 tilt-series each consisting of 9-14 tilt-slices, randomly selected SBs were investigated (L1a: 8 tilt series; 16 SBs with a total of 2094 SVs; L1b: 10 tilt series; 22 SBs with a total of 3740 SVs).

Quantitative analysis of the astrocytic coverage

Added in the revised version: Although methodologically associated with the possibility of a certain degree of error, astrocytic processes around synaptic complexes were mainly identified by their relatively clear cytoplasm compared to synaptic boutons or dendritic profiles; in addition by their irregular, stellate shape, by the presence of glycogen granules and bundles of intermediate filaments (Peters et al. 1991; Ventura and Harris 1999). With greater accuracy, astrocytic processes can be identified with immunolabeling, but this leads to poorer tissue preservation and resolution of cellular structures such as AZs. Among studies in rats (e.g. Rollenhagen et al. 2015; 2018) in a previous study on the human TLN (Yakoubi et al. 2019a) immunohistochemistry against glutamine synthetase, a key enzyme in astrocytes, was carried out. Since similar results were obtained without immunohistochemistry when identifying astrocytic processes around AZs, immunohistochemistry was omitted in the series of follow-up studies on the human TLN.

Earlier version: Astrocytic processes around synaptic complexes were mainly identified by their relatively clear cytoplasm compared to synaptic boutons or dendritic profiles; in addition by their irregular, stellate shape, by the presence of glycogen granules and bundles of intermediate filaments (Peters et al. 1991; Ventura and Harris 1999). To quantify the astrocytic coverage of synaptic complexes that may or may not constitute the ‘tripartite’ synapse in the human neocortex, the interactive software ImageJ (Schneider et al. 2012) was used. The first, the middle, and the last images of an individual EM series were used for a further quantitative volumetric analysis (Source Data 04, sheet 01). In each section of the same series used for the 3D-volume reconstructions, a grid (grid size 1 × 1 μm2) was placed over the EM image, and in each square, the abundance of fine astrocytic processes was documented throughout these images and averaged. Using the Cavalieri method [Unbiased Stereology: Three-Dimensional Measurement in Microscopy (Advanced Methods) Paperback-January 7, 2005, by Vyvyan Howard Matthew Reed], the (absolute) volume contribution of astrocytic processes was determined according to the Cavalieri estimator:

Revised version: To quantify the volume contribution of astrocytic processes in the human neocortex, the interactive software ImageJ (Schneider et al. 2012) was used. The first, the middle, and the last images of an individual EM series were used for a further quantitative volumetric analysis (Source Data 04, sheet 01). In each section of the same series used for the 3D-volume reconstructions, a grid (grid size 1 × 1 μm2) was placed over the EM image, and in each square, the abundance of fine astrocytic processes was documented throughout these images and averaged. Using the Cavalieri method [Unbiased Stereology: Three-Dimensional Measurement in Microscopy (Advanced Methods) Paperback-January 7, 2005, by Vyvyan Howard Matthew Reed], the (absolute) volume contribution of astrocytic processes was determined according to the Cavalieri estimator: where a(p) is the size of one square (0.8 × 0.8 μm2), P is the number of squares counted, and t is the thickness of the slice.

Removed in the revised version: where a(p) is the size of one square (0.8 × 0.8 μm2), P is the number of squares counted, and t is the thickness of the slice.

Earlier version: In addition, in consecutive sections spanning the complete AZ, the astrocytic coverage was evaluated under the following criteria: fine astrocytic processes were observed on 1) both sides of the AZ reaching as far as the synaptic cleft, 2) only one side of the AZ was covered by fine astrocytic processes and 3) AZs without astrocytic processes (Source Data 04, sheet 02). Values obtained for L1 either for SBs terminating on dendritic spines (n = 254) or shafts (n = 92) were compared to L4 (spines = 93; shafts = 41). Layer 4 was chosen because the volumetric analysis showed the highest portion of astrocytic processes in this layer compared to all other cortical layers. All values are given as percentages.

Revised version: In addition, in consecutive sections spanning the complete AZ, the astrocytic coverage was evaluated under the following criteria: fine astrocytic processes were observed on 1) both sides of the AZ reaching as far as the synaptic cleft, 2) only one side of the AZ was covered by fine astrocytic processes and 3) AZs without astrocytic processes (Source Data 04, sheet 02). Values obtained for L1 (L1a: n = 3 patients; L1b: 3 patients) either for SBs terminating on dendritic spines (L1a: n = 128; L1b: n=126) or shafts (L1a: n= 48; L1b: n = 44) were compared to L4 (n = 5 patients; dendritic spines n=93; dendritic shafts n=41). Layer 4 was chosen because the volumetric analysis showed the highest portion of astrocytic processes in this layer compared to all other cortical layers. All values are given as percentages.

Statistical analysis

Earlier version: The mean value ± standard deviation (SD), the median with the interquartile range (IQR), the coefficient of variation (CV), skewness, variance, and the coefficient of correlation (R2) were given for each structural parameter analyzed. The p-value was considered significant only if p < 0.05. Box- and Violin plots (Plotly 4.0.0 https://chart-studio.plotly.com) were generated to investigate inter-individual differences for each patient and structural parameter (Supplemental Figure 1, 2; Source Data 02).

Revised version: The mean value ± standard deviation (SD), the median with the interquartile range (IQR), the coefficient of variation (CV), skewness, variance, and the coefficient of correlation (R2) were given for each structural parameter analyzed. The p-value was considered significant only if p < 0.05. Box- and Violin plots (Plotly 4.0.0 https://chart-studio.plotly.com) were generated to investigate inter-individual differences for each patient and structural parameter (Supplemental Figures 1, 2; Source Data 02).

To test for significant differences, the non-parametric Kruskal-Wallis H-test with a subsequent Mann-Whitney u-test analysis was performed, using PAST 4.02 (Hammer et al. 2001). Correlation graphs between several structural parameters were then generated (Figure 4). The R2-values were interpreted as follows: 0, no linear correlation; 0-0.5, weak linear correlation; 0.5-0.8, good linear correlation; and 0.8-1.0, strong linear correlation. Furthermore, a freely available Fisher’s r-to-z-transformation calculator (Fisher’s z) was used to test for differences in R2 between L1a and L1b (p-value < 0.05).

Earlier version: Parts of the data acquisition were carried out in the context of a PhD Thesis (Sadeghi Dastjerdi A., 2024).Values obtained in the present study for L1 were compared partially to those obtained for L2, L3 (Girwert Ch., PhD Thesis in preparation), L4, L5 (Yakoubi et al. 2019a, b; Yakoubi R., PhD Thesis 2023) and L6 (Schmuhl-Giesen et al. 2022).

Revised version: For statistical comparison with other layers of the human TLN, L1a and L1b were grouped together as L1. Parts of the data acquisition were carried out in the context of a PhD Thesis (Sadeghi Dastjerdi A. 2024).Values obtained in the present study for L1 were compared partially to those obtained for L2, L3 (Girwert Ch., PhD Thesis in preparation), L4, L5 (Yakoubi et al. 2019a, b; Yakoubi R., PhD Thesis 2023) and L6 (Schmuhl-Giesen et al. 2022).